![CC-BY-SA](https://mirrors.creativecommons.org/presskit/buttons/88x31/svg/by-sa.svg)
This notebook was created by [Bernardo Freitas Paulo da Costa](http://www.im.ufrj.br/bernardofpc),
and is licensed under Creative Commons BY-SA

```python
import sys
```

```python
from inspect import getsourcelines
import re

def has_selfcall(fun):
    t,i = getsourcelines(fun)
    regexp = re.compile(fun.__name__)
    numcalls = sum([1 for l in t[i:] if regexp.search(l)])
    return numcalls > 0
```

# Instrumentar funções

Uma outra habilidade necessária no curso é você _instrumentar_ suas funções para extrair informação extra.

## Fibonacci

Implemente uma função recursiva `fibo` que calcula o $n$-ésimo número de Fibonacci.
Lembre-se que os números de Fibonacci são definidos por:
$$ \begin{cases} F_0 = 0 \\ F_1 = 1 \\ F _ {n+2} = F _ {n+1} + F_n \end{cases}. $$

<!-- grade_id: fibo, points: 0, grade: false, solution: true -->
```python
def fibo(n):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
fibo(10)
```

Testando parte da definição:

<!-- grade_id: fibo01, points: 1, grade: true, solution: false -->
```python
assert fibo(0) == 0
assert fibo(1) == 1
```

<!-- grade_id: fibo_rec_prop, points: 2, grade: true, solution: false -->
```python
assert fibo(15) == fibo(14) + fibo(13)
```

Testando que de fato ela é recursiva:

<!-- grade_id: fibo_selfcall, points: 1, grade: true, solution: false -->
```python
assert has_selfcall(fibo)
```

<!-- grade_id: fibo_rec, points: 1, grade: true, solution: false -->
```python
prevlim = sys.getrecursionlimit()
sys.setrecursionlimit(50)

try:
    fibo(50)
    sys.setrecursionlimit(prevlim)
    assert False
except RuntimeError as e:
    sys.setrecursionlimit(prevlim)
    assert e.args[0][:32] == 'maximum recursion depth exceeded'
```

Uma propriedade interessante dos números de Fibonacci é a divisibilidade:

<!-- grade_id: fibo_div, points: 1, grade: true, solution: false -->
```python
assert fibo(21) % fibo(7) == 0
```

## Contando

Agora, vamos instrumentar.

Escreva uma função `fibo_ncalls` que, além de retornar o $n$-ésimo número de Fibonacci,
também retorne quantas vezes ela foi chamada.

<!-- grade_id: fibo_ncalls, points: 0, grade: false, solution: true -->
```python
def fibo_ncalls(n):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
fibo_ncalls(3)
```

Vejamos que ela conta corretamente o número de Fibonacci:

<!-- grade_id: fibo_ncalls_ok, points: 2, grade: true, solution: false -->
```python
for x in range(10):
    assert fibo_ncalls(x)[0] == fibo(x)
```

E que `fibo_ncalls` continua sendo recursiva:

<!-- grade_id: fibo_ncalls_self, points: 1, grade: true, solution: false -->
```python
assert has_selfcall(fibo_ncalls)
```

<!-- grade_id: fibo_ncalls_rec, points: 1, grade: true, solution: false -->
```python
prevlim = sys.getrecursionlimit()
sys.setrecursionlimit(50)

try:
    fibo_ncalls(50)
    sys.setrecursionlimit(prevlim)
    assert False
except RuntimeError as e:
    sys.setrecursionlimit(prevlim)
    assert e.args[0][:32] == 'maximum recursion depth exceeded'
```

# Instrumentando a bisseção

Vamos fazer algo mais informativo do que apenas o número de vezes que a função se chama.
Vamos retornar _listas_ com informação extra.

## Retornando listas em funções recursivas

Para "esquentar os motores", vamos voltar à função fatorial:
implemente uma função **recursiva** `fatoriais(n)` que retorna a lista com os fatoriais de 0 até `n`.

Dica: como a chamada recursiva já dá praticamente a lista toda, o que você precisa fazer para incluir um elemento a mais no valor de retorno?

<!-- grade_id: fatoriais, points: 0, grade: false, solution: true -->
```python
def fatoriais(n):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
# Testes simples para você ver
fatoriais(0), fatoriais(4)
```

<!-- grade_id: fatoriais_list, points: 2, grade: true, solution: false -->
```python
# Retorna uma lista
assert isinstance(fatoriais(10), list)
assert len(fatoriais(100)) == 101
```

<!-- grade_id: fatoriais_4, points: 1, grade: true, solution: false -->
```python
# Dá certo num caso pequeno
assert fatoriais(4) == [1,1,2,6,24]
```

<!-- grade_id: fatoriais_sub, points: 1, grade: true, solution: false -->
```python
# Os fatoriais até 10 também estão na lista dos fatoriais até 13 ;-)
assert fatoriais(10) == fatoriais(13)[:11]
```

<!-- grade_id: fatoriais_div, points: 2, grade: true, solution: false -->
```python
l = fatoriais(100)
assert l[0] == 1
for i in range(1,101):
    assert i == l[i]//l[i-1]
```

## Agora, a bisseção

Escreva a bisseção de forma recursiva, e que retorne a lista de todos os intervalos considerados.
O primeiro intervalo é o próprio $[a,b]$, e o último é o que está "mais perto" da solução.

Vamos **representar** o intervalo $[a,b]$ pelo **par** `(a,b)`. Assim, a sua função retornará uma lista de pares:
    
    In: bissecao(cos, 0, 8, tol=1)
    Out: [(0, 8), (0, 4.0), (0, 2.0), (1.0, 2.0)]

<!-- grade_id: bissecao, points: 0, grade: false, solution: true -->
```python
def bissecao(f, a, b, tol=1e-6):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
from math import cos
bissecao(cos, 0, 8, tol=1)
```

<!-- grade_id: bissecao_list, points: 1, grade: true, solution: false -->
```python
assert isinstance(bissecao(cos, 0, 3), list)
```

<!-- grade_id: bissecao_prop, points: 2, grade: true, solution: false -->
```python
l = bissecao(cos, 0, 2)
x,y = l[-1]
assert abs(x-y) <= 1e-6
```

<!-- grade_id: bissecao_shrink, points: 2, grade: true, solution: false -->
```python
l = bissecao(cos, 0, 2)
fator = 2/1e-6
assert fator <= 2**(len(l)-1) < 2*fator
```

## Precisão relativa
Agora, modifique o teste de recursão para que a raiz seja encontrada com precisão **relativa** `tol`.
Cuidado com números negativos e casos onde o intervalo contém 0!

<!-- grade_id: bissecao_rel, points: 0, grade: false, solution: true -->
```python
def bissecao(f, a, b, tol=1e-6):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: biss_rel_type, points: 1, grade: true, solution: false -->
```python
assert isinstance(bissecao(cos, 0, 3), list)
```

```python
from math import pi
```

<!-- grade_id: biss_rel_err, points: 1, grade: true, solution: false -->
```python
l = bissecao(cos, 0, 2)
x,y = l[-1]
assert abs(x-y)/(pi/2) <= 1e-6
```

<!-- grade_id: biss_rel_hard, points: 2, grade: true, solution: false -->
```python
l = bissecao(cos, -0.87, -3.43, tol=1e-1)
x,y = l[-1]
assert abs(x-y)/(pi/2) <= 1e-1
```

<!-- grade_id: biss_rel_negs, points: 1, grade: true, solution: false -->
```python
l = bissecao(cos, -5,10)
x,y = l[-1]
assert (x - pi/2) <= 1e-6 * pi/2
assert (y - pi/2) <= 1e-6 * pi/2
```

<!-- grade_id: biss_rel_precise_err, points: 2, grade: true, solution: false -->
```python
l = bissecao(cos, -3,7)
x,y = l[-1]
abs_err = abs(y-x)
assert 0.5e-6 < abs_err/(3*pi/2) <= 1e-6
assert 10 * 2**(1-len(l)) == abs_err
```

<!-- grade_id: biss_rel_shrink, points: 2, grade: true, solution: false -->
```python
# Verificando que os intervalos de fato são divididos por 2 a cada vez
l = bissecao(cos, -3,7, 1e-14)
assert l[0] == (-3,7)
diam = 10
for x,y in l[1:]:
    # O que vem depois da vírgula do assert aparece no caso de dar erro
    assert abs(y-x) == diam/2, ((y-x), diam/2)
    diam = y-x
```