![CC-BY-SA](https://mirrors.creativecommons.org/presskit/buttons/88x31/svg/by-sa.svg)
This notebook was created by [Bernardo Freitas Paulo da Costa](http://www.im.ufrj.br/bernardofpc),
and is licensed under Creative Commons BY-SA

```python
import numpy as np
import matplotlib.pyplot as plt
```

Inclua abaixo seus métodos numéricos para encontrar raízes.

<!-- grade_id: cell-99416d31c0cf2a14, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

# Encontrando raízes de polinômios

O objetivo desta prova é desenvolver um método eficiente para encontrar
todas as raízes reais de um polinômio com coeficientes reais.

# Parte 1: polinômios de grau pequeno (independente das partes 2 e 3)

## 1.1 Funções afins

Suponha que $f$ é um polinômio de grau 1.  Escreva uma função que retorna a sua raiz.

<!-- grade_id: raiz1, points: 0, grade: false, solution: true -->
```python
def raiz1(f):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
def f_lin(x):
    return 4*x + 1.342

raiz1(f_lin)
```

```python
f_lin(raiz1(f_lin))
```

<!-- grade_id: raiz1_test, points: 2, grade: true, solution: false -->
```python
def f(x):
    return -423.33*x + 784

assert raiz1(f) == 1.851983086480996
assert f(raiz1(f)) == 0
```

<!-- grade_id: raiz1_random, points: 2, grade: true, solution: false -->
```python
np.random.seed(1)
# Gerando 100 polinômios aleatórios de grau 1
for a,b in np.random.randn(100, 2):
    def f(x):
        return a*x + b
    r = raiz1(f)
    assert np.abs( f(r) ) < 2e-14 * (np.abs(r) + np.abs(b)), (a,b,r,f(r))
```

Agora, escreva uma função que diz se $f$ é crescente ou decrescente,
retornando 1 ou -1 respectivamente.

<!-- grade_id: crescente1, points: 0, grade: false, solution: true -->
```python
def crescente1(f):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: crescente1_test, points: 1, grade: true, solution: false -->
```python
assert crescente1(f_lin) == 1
```

<!-- grade_id: crescente1_random, points: 1, grade: true, solution: false -->
```python
np.random.seed(2)
for a,b in np.random.randn(100, 2):
    def f(x):
        return a*x + b
    r = crescente1(f)
    assert r * a > 0
```

## 1.2 Funções quadráticas

Suponha agora que $f$ é um polinômio de grau 2.

### 1.2.1 Concavidade

Escreva uma função que descobre se a concavidade de $f$ está para cima ou para baixo.

<!-- grade_id: concavidade, points: 0, grade: false, solution: true -->
```python
def concavidade(f):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
def f_quad(x):
    return x**2 + 6*x - 10

concavidade(f_quad)
```

<!-- grade_id: concavidate_test, points: 1, grade: true, solution: false -->
```python
assert concavidade(f_quad) == 1
```

<!-- grade_id: concavidade_random, points: 2, grade: true, solution: false -->
```python
np.random.seed(3)
for a,b,c in np.random.randn(100, 3):
    def f(x):
        return a*x**2 + b*x + c
    ans = concavidade(f)
    assert ans == np.sign(a), (ans, a)
```

### 1.2.2 Vértice

Agora, faça uma função que encontra o vértice da parábola determinada por $f$.

<!-- grade_id: vertice, points: 0, grade: false, solution: true -->
```python
def vertice(f):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
vertice(f_quad)
```

<!-- grade_id: vertice_test, points: 1, grade: true, solution: false -->
```python
assert vertice(f_quad) == -3
```

<!-- grade_id: vertice_random, points: 2, grade: true, solution: false -->
```python
np.random.seed(4)
for a,b,c in np.random.randn(100, 3):
    def f(x):
        return a*x**2 + b*x + c
    ans = vertice(f)
    assert np.abs(ans + b/2/a) < 1e-13 * (np.abs(a) + np.abs(b) + np.abs(ans)), (a,b,-b/2/a, ans+b/2/a)
```

### 1.2.3 Raízes

Agora, faça uma função que encontra as duas raízes de $f$, ou que diz que a função não tem raiz.
Nesta prova, vamos supor que não há raízes duplas, para simplificar.

(Esta função é relativamente fácil se você calculou o vértice)

<!-- grade_id: raiz2, points: 0, grade: false, solution: true -->
```python
def raiz2(f):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
raiz2(f_quad)
```

<!-- grade_id: raiz2_test, points: 2, grade: true, solution: false -->
```python
r1, r2 = raiz2(f_quad)
assert r1 <= r2
assert np.allclose( [f_quad(r1), f_quad(r2)], [0,0], atol=1e-13, rtol=1e-13 )
```

<!-- grade_id: raiz2_none, points: 2, grade: true, solution: false -->
```python
def f_quad2(x): return x**2 + 0.1
assert raiz2(f_quad2) is None
```

## 1.3 Contando as raizes de uma cúbica

Suponha que você tem tanto $f$, um polinômio de grau 3, como $f'$, um polinômio de grau 2.

Escreva uma função que diga quantas raízes o polinômio tem,
olhando para as raízes de $f'$.

<!-- grade_id: nroots, points: 0, grade: false, solution: true -->
```python
def nroots(f, df):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: nroots_test1, points: 1, grade: true, solution: false -->
```python
def f_cubic(x): return x**3 + x + 1
def df_cubic(x): return 3*x**2 + 1

assert nroots(f_cubic, df_cubic) == 1
```

```python
xs = np.arange(-2,2,0.01)
plt.plot(xs, f_cubic(xs))
plt.axhline(linestyle='--', color='k', linewidth=1)
plt.show()
```

<!-- grade_id: nroots_test3, points: 1, grade: true, solution: false -->
```python
def f_cubic(x): return x**3 - 3*x + 1
def df_cubic(x): return 3*x**2 - 3

assert nroots(f_cubic, df_cubic) == 3
```

```python
xs = np.arange(-2,2,0.01)
plt.plot(xs, f_cubic(xs))
plt.axhline(linestyle='--', color='k', linewidth=1)
plt.show()
```

<!-- grade_id: nroots_random, points: 2, grade: true, solution: false -->
```python
# Aqui há testes escondidos
```

# Parte 2: Álgebra computacional

Vamos representar um polinômio $p(x) = a_0 + a_1 x + \ldots + a_d x^d$ pela lista de seus coeficientes, nesta ordem:
`[a0, a1, ..., ad]`.

## 2.1: Calculando polinômios

Escreva uma função que, dados os coeficientes, e um número real $x$, retorna $p(x)$.

(Não se preocupe em fazer a função ser vetorizada em `x`, não vamos usar muito)

<!-- grade_id: evalp, points: 0, grade: false, solution: true -->
```python
def eval_p(p, x):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
eval_p([1,2,3], 4)
```

<!-- grade_id: evalp_test, points: 2, grade: true, solution: false -->
```python
assert eval_p([1,2,3], 4) == 57
```

Agora, escreva uma função análoga, que retorna a _função polinomial_ determinada pelos coeficientes.

<!-- grade_id: funp, points: 0, grade: false, solution: true -->
```python
def fun_p(p):
    # YOUR CODE HERE
    raise NotImplementedError()
```

## 2.2: Derivada

Escreva uma função que, dados os coeficientes do polinômio, retorna os coeficientes de sua derivada.

<!-- grade_id: dpol, points: 0, grade: false, solution: true -->
```python
def d_pol(p):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
d_pol([3,1,5,7,2])
```

<!-- grade_id: dpol_test, points: 2, grade: true, solution: false -->
```python
assert all(d_pol([3,1,5,7,2]) == np.array([ 1, 10, 21,  8]))
```

Faça agora um gráfico de $x^5 - 4x + 1$ e de sua derivada, na mesma figura, usando `eval_p` ou `fun_p`.

<!-- grade_id: dpol_graph, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## 2.3 Sinal no infinito

Escreva uma função que retorna o sinal de $p(x)$ em $\pm \infty$, dados os coeficientes.

<!-- grade_id: sign_inf, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Retirar alguns desses

<!-- grade_id: sign_inf_test1, points: 1, grade: true, solution: false -->
```python
assert sign_inf([1,2,3,4,5,6]) == 1
assert sign_inf([1,2,3,4,5,6], minus=True) == -1
```

<!-- grade_id: sign_inf_test2, points: 1, grade: true, solution: false -->
```python
assert sign_inf([1,2,3,4,5,6,7]) == 1
assert sign_inf([1,2,3,4,5,6,7], minus=True) == 1
```

<!-- grade_id: sign_inf_test3, points: 1, grade: true, solution: false -->
```python
assert sign_inf([1,2,3,4,5,-6]) == -1
assert sign_inf([1,2,3,4,5,-6], minus=True) == 1
```

<!-- grade_id: sign_inf_test4, points: 1, grade: true, solution: false -->
```python
assert sign_inf([1,2,3,4,5,6,-7]) == -1
assert sign_inf([1,2,3,4,5,6,-7], minus=True) == -1
```

## 2.4 Raízes de polinômios do segundo grau (independente da parte 3)

Escreva uma função que, dados os coeficientes `[c,b,a]` de um polinômio do segundo grau,
retorna as duas raízes reias, ou `None` se não houver.

<!-- grade_id: raiz2c, points: 0, grade: false, solution: true -->
```python
def raiz2_c(p):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
raiz2_c([-1,-2,3])
```

<!-- grade_id: raiz2c_test, points: 1, grade: true, solution: false -->
```python
r1, r2 = raiz2_c([-1,-2,3])
assert r1 == -1/3
assert r2 == 1
```

<!-- grade_id: raiz2c_random, points: 2, grade: true, solution: false -->
```python
np.random.seed(5)
for a,b,c in np.random.randn(100, 3):
    def f(x):
        return a*x**2 + b*x + c
    raizes = raiz2(f)
    if raizes is None:
        assert raiz2_c([c,b,a]) is None
    else:
        raizes_c = raiz2_c([c,b,a])
        assert np.allclose(raizes, raizes_c, rtol=1e-12, atol=1e-12), (raizes, raizes_c)
```

## 2.5 Coeficientes do polinômio de segundo grau

Escreva uma função tal que, dados dois números reais $r_1$ e $r_2$
retorna uma lista com os coeficientes `[c,b,a]`
tal que $r_1$ e $r_2$ sejam as raízes do polinômio do segundo grau correspondente.

<!-- grade_id: makecoefs, points: 0, grade: false, solution: true -->
```python
def make_coefs(r1, r2):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
make_coefs(1,2)
```

<!-- grade_id: makecoefs_random, points: 2, grade: true, solution: false -->
```python
for r1,r2 in np.random.randn(100,2):
    p = make_coefs(r1,r2)
    assert abs(eval_p(p,r1)) < 1e-12
    assert abs(eval_p(p,r2)) < 1e-12
```

## 2.6 Erros

Seja $P_\varepsilon$ o polinômio cujas raízes são $1 \pm \varepsilon$.
Calcule a menor raíz de $P_\varepsilon$ pela função `raiz2_c()`,
e observe como o erro evolui para diversos valores de $\varepsilon$.

<!-- grade_id: errors_graph, points: 3, grade: true, solution: true -->
```python
eps = np.logspace(-40,5,base=2, num=500)
# YOUR CODE HERE
raise NotImplementedError()
```

Você consegue explicar a origem deste erro?

<!-- grade_id: errors_explain, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

# Parte 3: Polinômios de grau arbitrário

Para encontrar raízes de um polinômio $p$, vamos olhar para as raízes da sua derivada,
e com isso obter todos os intervalos possíveis onde há raízes de $p$.

## 3.1 Raízes encaixadas

Seja $f$ uma função diferenciável, e $f'$ sua derivada.
**Prove** (no notebook ou numa folha avulsa) que entre duas raizes consecutivas de $f'$,
existe no máximo uma raiz de $f$.

(Se você precisar, pode supor que $f'$ é contínua)

<!-- grade_id: nested_roots, points: 4, grade: true, solution: true -->
YOUR ANSWER HERE

## 3.2 Raízes até o infinito

Além das possíveis raízes de $p$ que haja entre as raízes de $p'$,
pode haver raízes "para além das extremas", em cada direção.

Dados os coeficientes de um polinômio, e um ponto $x$,
retorne um ponto $X$ tal que $p(X)$ tenha o mesmo sinal que $p(+\infty)$,
e que seja maior ou igual a $x$.

<!-- grade_id: aux1, points: 0, grade: false, solution: true -->
```python
def aux1(p, x):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: aux1_random, points: 2, grade: true, solution: false -->
```python
np.random.seed(8)
for p in np.random.randn(100, 10):
    x = aux1(p,0)
    assert eval_p(p,x) * sign_inf(p) > 0
    assert x > 0
```

Faça uma função análoga para encontrar um ponto menor do que $x$, cujo sinal seja o mesmo que o sinal em $-\infty$.

<!-- grade_id: aux2, points: 0, grade: false, solution: true -->
```python
def aux2(p, x):
    # YOUR CODE HERE
    raise NotImplementedError()
```

## 3.3 O passo de indução

Escreva uma função que, dada uma lista de pontos `[x0, x1, ..., xn]`,
encontra um zero em cada intervalo, se houver,
incluindo os intervalos infinitos $(-\infty,x_0)$ e $(x_n, \infty)$.

<!-- grade_id: inducao, points: 0, grade: false, solution: true -->
```python
def raizes_intervalos(p, zeros_der):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: inducao_test1, points: 2, grade: true, solution: false -->
```python
p = [0,-2,1]
zeros_der = [1]
ans = [0,2]
assert np.allclose(ans, raizes_intervalos(p, zeros_der), atol=1e-13, rtol=1e-13)
```

<!-- grade_id: inducao_test2, points: 2, grade: true, solution: false -->
```python
p = [2,-2,1]
zeros_der = [-1]
assert raizes_intervalos(p, zeros_der) == []
```

```python
p = [3,5,-20,4,1]
zeros_der = [-5.035286135110316, 0.13031598251839685, 1.9049701525919194]
raizes_intervalos(p, zeros_der)
```

<!-- grade_id: inducao_test3, points: 2, grade: true, solution: false -->
```python
p = [3,5,-20,4,1]
zeros_der = [-5.035286135110316, 0.13031598251839685, 1.9049701525919194]
ans = [-6.9654801153285089,
 -0.27707634751771126,
 0.58488321666365461,
 2.657673246182565]
assert np.allclose(ans, raizes_intervalos(p, zeros_der), atol=1e-13, rtol=1e-13)
```

<!-- grade_id: inducao_test4, points: 2, grade: true, solution: false -->
```python
p = [40,5,-20,4,1]
r = raizes_intervalos(p, zeros_der)
assert len(r) == 2
assert np.abs(eval_p(p, r[0])) < 1e-12
assert np.abs(eval_p(p, r[1])) < 1e-12
```

## 3.4 Juntando tudo

Agora, escreva uma função (recursiva!) que calcula todas as raízes reais de um polinômio.
Não se esqueça
- de retornar as raízes em ordem (para poder usar `raizes_intervalos()` recursivamente)
- de ter um caso base para grau 1!

<!-- grade_id: allroots, points: 0, grade: false, solution: true -->
```python
def all_roots(p):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
coefs = [0.2,-2,1,1]
all_roots(coefs)
```

<!-- grade_id: allroots_test1, points: 2, grade: true, solution: false -->
```python
coefs = [0.2,-2,1,1]
ans = [-2.0324501292055701, 0.10624345669581943, 0.92620667250975075]
assert np.allclose(all_roots(coefs), ans, atol=1e-13, rtol=1e-13)
```

<!-- grade_id: allroots_test2, points: 2, grade: true, solution: false -->
```python
coefs = [0.2,-2,1,1]
for r in all_roots(coefs):
    assert np.abs(eval_p(coefs,r)) < 1e-13
```

```python
coefs = [0.2,10,1,1,-2,-3,1]
all_roots(coefs)
```

<!-- grade_id: allroots_test3, points: 2, grade: true, solution: false -->
```python
for r in all_roots(coefs): 
    assert np.abs(eval_p(coefs,r)) < 1e-12
```