![CC-BY-SA](https://mirrors.creativecommons.org/presskit/buttons/88x31/svg/by-sa.svg)
This notebook was created by [Bernardo Freitas Paulo da Costa](http://www.im.ufrj.br/bernardofpc),
and is licensed under Creative Commons BY-SA

# Tempo de convergência

Vamos neste teste observar quanto tempo o Método de Newton leva para convergir!

```python
import numpy as np
```

## Parte 0:

Newton contando iterações:
modifique o método de Newton para que ele retorne um par com a raíz e o número de iterações.

<!-- grade_id: newton_nsteps, points: 0, grade: false, solution: true -->
```python
def newton_nsteps(f, df, x, prec=1e-12, maxsteps=100):
    # YOUR CODE HERE
    raise NotImplementedError()
```

Testando que tanto o valor de retorno como o número de passos estão "contando do lugar certo".
Se você estiver obtendo erros de "contar" um a mais / a menos, lembre que o método de Newton:
- pára quando o passo esteja menor do que `prec`
- mas **anda** este último passo, já que seria desperdício não usar um valor já calculado,
- e retorna o número de passos efetuados.

<!-- grade_id: newton_sin1, points: 2, grade: true, solution: false -->
```python
assert newton_nsteps(np.sin, np.cos, 1) == (0.0, 5)
```

<!-- grade_id: newton_sin1000, points: 2, grade: true, solution: false -->
```python
ans, n = newton_nsteps(np.sin, np.cos, 1000)
assert n == 5
assert np.sin(ans) == -1.7627486140082825e-14
```

## Parte 1: Uma inversa simples

Escreva uma função `arcoseno`, que calcule pelo método de Newton o arco seno de um número entre -1 e 1,
e também retorne o número de iterações necessárias para isso.
Como ainda temos um parâmetro "sobrando", vou pedir para vocês começarem sempre em $x = 0$.

<!-- grade_id: arcoseno, points: 0, grade: false, solution: true -->
```python
def arcoseno(y):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
arcoseno(0.5)
```

### Testes básicos

<!-- grade_id: asin05, points: 2, grade: true, solution: false -->
```python
ans, n = arcoseno(0.5)
assert ans == np.arcsin(0.5)
assert n == 5
```

<!-- grade_id: asin08, points: 1, grade: true, solution: false -->
```python
ans, n = arcoseno(0.8)
assert abs(ans - np.arcsin(0.8)) < 2e-16
assert n == 6
```

<!-- grade_id: asin099, points: 0, grade: true, solution: false -->
```python
ans, n = arcoseno(0.99)
assert abs(ans - np.arcsin(0.99)) < 4e-16
assert n == 8
```

### Testando propriedades

Uma propriedade bem razoável do método:

<!-- grade_id: asin_prop1, points: 2, grade: true, solution: false -->
```python
for x in np.random.rand(5):
    a,m = arcoseno(x)
    b,n = arcoseno(-x)
    assert (a == -b)
    assert (m == n)
```

Como você descreveria em palavras esta propriedade?

Uma outra propriedade:

<!-- grade_id: asin_prop2, points: 2, grade: true, solution: false -->
```python
prevans, prevn = 0, 0
for x in np.arange(0.1,1,0.03):
    ans, n = arcoseno(x)
    assert prevans < ans
    assert prevn <= n
    prevans, prevn = ans, n
```

### Observando um pouco mais:

Eis aqui um resultado para pensar:
mesmo que a tolerância sugerida seja `1e-12`, o erro foi quase `1e-8`.
E levamos 28 passos.

```python
ans, n = arcoseno(1)
np.pi/2 - ans, n
```

## Parte 2: Um arcoseno melhor

Sabemos que $\sin(x) \sim x$, logo também deve ser verdade que $x \sim \arcsin(x)$.
Use isso para ter um ponto inicial melhor para o método de Newton:

<!-- grade_id: asin_fast, points: 0, grade: false, solution: true -->
```python
def arcoseno_rapido(y):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: fast_asin05, points: 1, grade: true, solution: false -->
```python
ans, n = arcoseno_rapido(0.5)
assert ans == np.arcsin(0.5)
assert n == 4
```

<!-- grade_id: fast_asin099, points: 1, grade: true, solution: false -->
```python
ans, n = arcoseno_rapido(0.99)
assert abs(ans - np.arcsin(0.99)) < 4e-16
assert n == 7
```

Veja só:

<!-- grade_id: asin_fastasin, points: 2, grade: true, solution: false -->
```python
for x in np.arange(0.1,1,0.03):
    a, m = arcoseno(x)
    b, n = arcoseno_rapido(x)
    assert a == b
    assert m == n + 1
```

Porque isso acontece sempre?
Ajudou muito ter uma aproximação muito boa do arco seno?

## Parte 3: Um gráfico e uma bisseção

Faça um gráfico do número de iterações necessárias para calcular o arco seno dos números de -0.9 a 0.9.
Depois, faça outro com o número de iterações para 0.9 até 0.999.

```python
import matplotlib.pyplot as plt
```

<!-- grade_id: graf_niter, points: 3, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Agora, olhando o gráfico para ter os chutes iniciais,
faça uma "bisseção" para descobrir a partir de qual número real
a sua função `arcoseno` necessita de 6 iterações para convergir.

<!-- grade_id: niter6, points: 0, grade: false, solution: true -->
```python
# Dê abaixo a resposta na forma
# x = ...
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: test_niter6, points: 3, grade: true, solution: false -->
```python
# Testando que de fato usamos 6 iterações
_, n = arcoseno(x)
assert n == 6
# e que o número anterior (no computador) leva 5
_, n = arcoseno(x - 1e-16)
assert n == 5
# Ah, porque subtrair 1e-16 garante que será o anterior??
```

Agora, vamos ver quanto tempo demora para chegar em 20 iterações!

Sugestão: implemente a função abaixo, vai ajudar para a última questão também!

<!-- grade_id: cell-e9a69acc43191e49, points: 0, grade: false, solution: true -->
```python
def biss_newton_n(a,b, n):
    """Encontra o ponto $x \in [a,b]$ tal que `arcoseno(x)` leva $n$ iterações, e o "anterior" leva $n-1$. """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: niter_20, points: 0, grade: false, solution: true -->
```python
# Dê abaixo a resposta na forma
# x = ...
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: test_niter20, points: 2, grade: true, solution: false -->
```python
_, n = arcoseno(x)
assert n == 20
_, n = arcoseno(x - 1e-16)
assert n == 19
```

Enfim, faça um gráfico dos pontos $x$ que são as mudanças de número de iteração, para vários valores de $n$

<!-- grade_id: graf_niters, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```