![CC-BY-SA](https://mirrors.creativecommons.org/presskit/buttons/88x31/svg/by-sa.svg)
This notebook was created by [Bernardo Freitas Paulo da Costa](http://www.im.ufrj.br/bernardofpc),
and is licensed under Creative Commons BY-SA

```python
import numpy as np
import matplotlib.pyplot as plt
```

# 1. A _Regula Falsi_

O método da _Regula Falsi_ é uma mistura entre o método da bisseção e o da secante:
- Tomamos dois pontos $a$ e $b$ tais que $f(a) \cdot f(b) < 0$, como no método da bisseção;
- Usamos como próximo ponto a aproximação secante $z = \frac{a f(b) - b f(a)}{f(b) - f(a)}$;
- E continuamos no intervalo $[a,z]$ ou $[z,b]$ em que $f$ mude de sinal, como na bisseção.

## 1.1 Código

Implemente esta variante.  Vamos usar tanto `prec` (da bisseção) como `maxiter` (da secante) como critérios de parada:
- se o intervalo $(a,b)$ tiver comprimento menor do que `prec`, diremos que o método convergiu, e retornamos o intervalo atual;
- se não for possível fazer mais iterações, (também) retornamos o intervalo atual;
- senão, calculamos o novo ponto, escolhemos o novo intervalo, e fazemos a recursão.

<!-- grade_id: rf, points: 0, grade: false, solution: true -->
```python
def regula_falsi(f,a,b,prec=1e-8,maxiter=100):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
import funcoes_teste as ft
```

<!-- grade_id: rf_1, points: 2, grade: true, solution: false -->
```python
assert np.allclose(regula_falsi(ft.g1, 0, 2), (0.69314718055994529, 2), rtol=1e-15, atol=1e-15)
```

<!-- grade_id: rf_2, points: 1, grade: true, solution: false -->
```python
assert np.allclose(regula_falsi(ft.g2, 0, 2), (1.414213562373095, 1.4142135623730951), rtol=1e-15, atol=1e-15)
```

<!-- grade_id: rf_3, points: 1, grade: true, solution: false -->
```python
assert np.allclose(regula_falsi(ft.g3, 0, 2), (0.84936886239267306, 0.86510529519350243), rtol=1e-15, atol=1e-15)
```

## 1.2 Instrumentando o código

Agora, modifique o método para que ele retorne a lista de todos os intervalos gerados.

<!-- grade_id: rf_list, points: 0, grade: false, solution: true -->
```python
def regula_falsi_list(f,a,b,prec=1e-8,maxiter=100):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: rf_1_101, points: 2, grade: true, solution: false -->
```python
assert len(regula_falsi_list(ft.g1, 0, 2)) == 101
```

<!-- grade_id: rf_2_23, points: 2, grade: true, solution: false -->
```python
assert len(regula_falsi_list(ft.g2, 0, 2)) == 23
```

<!-- grade_id: rf_3_101, points: 1, grade: true, solution: false -->
```python
assert len(regula_falsi_list(ft.g3, 0, 2)) == 101
```

## 1.3 Convergência

Faça uma função que, dada uma lista de intervalos (como a que é retornada pela `regula_falsi_list`!),
faz o gráfico do comprimento dos intervalos.

<!-- grade_id: plot_interval_lens, points: 2, grade: true, solution: true -->
```python
def plot_interval_lens(l):
    # YOUR CODE HERE
    raise NotImplementedError()
```

Agora, verifique como evolui o comprimento dos intervalos, para cada uma das três funções acima.

```python
plt.figure(figsize=(15,4))

plt.subplot(1,3,1)
plot_interval_lens(regula_falsi_list(ft.g1, 0, 2))

plt.subplot(1,3,2)
plot_interval_lens(regula_falsi_list(ft.g2, 0, 2))

plt.subplot(1,3,3)
plot_interval_lens(regula_falsi_list(ft.g3, 0, 2))
plt.yscale('log')

plt.show()
```

Isso é coerente com o comportamento da função `regula_falsi_list`?

<!-- grade_id: comment_lens, points: 3, grade: true, solution: true -->
YOUR ANSWER HERE

## 1.4 Convergência separada

Agora, faça uma função que, dada a mesma lista de intervalos e um número real `ans`,
faz o gráfico da distância de cada uma das extremidades a `ans`.
Não esqueça de indicar qual é a extremidade correspondente a cada curva, num `label` para a legenda ;-)

<!-- grade_id: plot_interval_dist, points: 3, grade: true, solution: true -->
```python
def plot_interval_distance(l, ans):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
plot_interval_distance(regula_falsi_list(ft.g1,0,2), np.log(2))
plt.legend()
plt.show()
```

O que aconteceu com o método?

<!-- grade_id: comments_dist_g1, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

Faça o gráfico análogo para $g_2$, e comente.

<!-- grade_id: graph_dist_g2, points: 1, grade: true, solution: true -->
```python
# Faça contas aqui, comente na caixa abaixo!
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: comment_dist_g2, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

## 2 Modificando a _Regula Falsi_

O que vimos antes mostra que a _Regula Falsi_ possui alguns defeitos.
Vamos explorar duas modificações:
- alterar o critério de parada
- alterar o método de seleção do intervalo

## 2.1 Critério de parada

Em vez de pararmos quando o intervalo "de possibilidade da raiz" for pequeno
(o que é bom para a garantia de erro, como na bisseção),
vamos parar quando o "passo" for pequeno.
Ou seja, se a extremidade que "se mexer" for "se mexer muito pouco" (ou seja, menos que `prec`), paramos.
Note que este critério só pode ser calculado após determinar o "novo ponto".

<!-- grade_id: rf_new, points: 0, grade: false, solution: true -->
```python
def regula_falsi_list_new(f,a,b,prec=1e-8,maxiter=100):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: rf_new_1, points: 2, grade: true, solution: false -->
```python
l = regula_falsi_list_new(ft.g1, 0,2)
assert len(l) == 29
assert np.allclose(l[-1], (np.log(2),2), rtol=1e-8, atol=1e-8)
```

<!-- grade_id: rf_new_2, points: 1, grade: true, solution: false -->
```python
l = regula_falsi_list_new(ft.g2, 0,2)
assert len(l) == 13
assert np.allclose(l[-1], (np.sqrt(2),2), rtol=1e-8, atol=1e-8)
```

<!-- grade_id: rf_new_3, points: 1, grade: true, solution: false -->
```python
l = regula_falsi_list_new(ft.g3, 0,2)
assert len(l) == 8
assert np.abs(ft.g3(l[-1][0])) < 1e-8
```

O que acontece se quisermos esperar o passo da _Regula Falsi_ ficar menor do que `1e-12`?
Faça alguns experimentos numéricos, e depois comente.

<!-- grade_id: num_xp, points: 1, grade: true, solution: true -->
```python
# Faça contas aqui, comente na caixa abaixo!
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: num_xp_comments, points: 3, grade: true, solution: true -->
YOUR ANSWER HERE

## 2.2 Passos combinados

Uma outra tática é combinar a bisseção com a _Regula Falsi_.
Assim, teremos passos "alternados", uma vez com o ponto médio, outra vez com o ponto da secante.
Vamos fazer isso com um par de funções que se chamam mutuamente (isso é chamado de "co-rotinas").

Como vamos usar bisseção, não precisamos mais usar `maxiter`.
Por outro lado, existe agora um "risco" de acertar a raiz em cheio,
então adicione também um teste para $f(z) = 0$.

<!-- grade_id: rf_biss1, points: 0, grade: false, solution: true -->
```python
def rf_biss_list(f, a,b, prec=1e-8):
    """RF + bisseção, passo da Regula Falsi. """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: rf_biss2, points: 0, grade: false, solution: true -->
```python
def biss_rf_list(f, a,b, prec=1e-8):
    """bisseção + RF, passo da bisseção. """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: rf_biss_1, points: 2, grade: true, solution: false -->
```python
l = rf_biss_list(ft.g1, 0, 2)
assert len(l) == 19
assert ft.g1(l[-1][0]) == 0
```

<!-- grade_id: rf_biss_2, points: 1, grade: true, solution: false -->
```python
l = rf_biss_list(ft.g2, 0, 2)
assert len(l) == 16
assert abs(l[-1][1] - l[-1][0]) < 1e-8
```

<!-- grade_id: rf_biss_3, points: 1, grade: true, solution: false -->
```python
l = rf_biss_list(ft.g3, 0, 2)
assert len(l) == 23
assert ft.g3(l[-1][0]) == 0
```

## 2.3 Comparando

Faça um gráfico da distância à raiz, usando a _Regula Falsi_ original, e esta.
Sugestão: modifique a sua função para ter um `label` a mais em argumento.

<!-- grade_id: plot_interval_label, points: 0, grade: false, solution: true -->
```python
def plot_interval_distance(l, ans, label=''):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: rf_vs_rf_biss, points: 3, grade: true, solution: true -->
```python
# Faça o gráfico aqui!
# YOUR CODE HERE
raise NotImplementedError()
```

Explique porque esta mudança
1. Faz sentido
2. De fato acelerou a convergência nos casos acima

<!-- grade_id: comment_rf_biss, points: 3, grade: true, solution: true -->
YOUR ANSWER HERE