![CC-BY-SA](https://mirrors.creativecommons.org/presskit/buttons/88x31/svg/by-sa.svg)
This notebook was created by [Bernardo Freitas Paulo da Costa](http://www.im.ufrj.br/bernardofpc),
and is licensed under Creative Commons BY-SA

# Listas e funções.


Uma habilidade fundamental deste curso é se sentir à vontade para manipular listas e funções.
Nada melhor que um exercício que misture as duas!

```python
from math import sin, cos
```

```python
from inspect import getsourcelines
import re

def haslistcomp(fun):
    t,_ = getsourcelines(fun)
    regexp = re.compile('\[.* for .* in .*\]')
    for l in t:
        if regexp.search(l):
            return True
    return False

def has_selfcall(fun):
    t,i = getsourcelines(fun)
    regexp = re.compile(fun.__name__)
    numcalls = sum([1 for l in t[i:] if regexp.search(l)])
    return numcalls > 0
```

# Aplicando funções em listas

Escreva uma função `aplicar_na_lista` que recebe uma função `f` e uma lista `l`
e retorna a lista resultante de aplicar `f` em cada um dos elementos de `l`.

Use uma [_list comprehension_](https://docs.python.org/3/tutorial/datastructures.html#list-comprehensions) para fazer isso.

<!-- grade_id: aplicar_na_lista, points: 0, grade: false, solution: true -->
```python
def aplicar_na_lista(f,l):
    # YOUR CODE HERE
    raise NotImplementedError()
```

Vamos ver que está funcionando:

```python
# Criando uma lista
l = list(range(10))
l
```

<!-- grade_id: aplicar_na_lista1, points: 1, grade: true, solution: false -->
```python
assert aplicar_na_lista(cos, l) == [1.0,
 0.5403023058681398,
 -0.4161468365471424,
 -0.9899924966004454,
 -0.6536436208636119,
 0.28366218546322625,
 0.960170286650366,
 0.7539022543433046,
 -0.14550003380861354,
 -0.9111302618846769]
```

```python
def odd(x):
    return x%2

aplicar_na_lista(odd, l)
```

<!-- grade_id: aplicar_na_lista2, points: 1, grade: true, solution: false -->
```python
assert aplicar_na_lista(odd, l) == [0, 1, 0, 1, 0, 1, 0, 1, 0, 1]
```

Agora, vamos ver se de fato você usou uma list comprehension:

<!-- grade_id: aplicar_na_lista_listcomp, points: 3, grade: true, solution: false -->
```python
assert haslistcomp(aplicar_na_lista)
```

Veja só que propriedade interessante!

```python
from numpy.random import randint
```

```python
n_el = randint(100)
x = list(range(n_el*2))
sum(aplicar_na_lista(odd, x)), n_el
```

Podemos usar propriedades interessantes para testar a sua função de forma aleatória!

<!-- grade_id: aplicar_na_lista_rand, points: 3, grade: true, solution: false -->
```python
for _ in range(10):
    n_el = randint(100)
    x = list(range(n_el*2))
    assert sum(aplicar_na_lista(odd, x)) == n_el
```

Uma propriedade mais geral ainda:

<!-- grade_id: aplicar_na_lista_concat, points: 3, grade: true, solution: false -->
```python
l2 = list(range(50,60))
assert aplicar_na_lista(odd, l+l2) == aplicar_na_lista(odd, l) + aplicar_na_lista(odd, l2)
```

<!-- grade_id: aplicar_na_lista_mult, points: 2, grade: true, solution: false -->
```python
assert aplicar_na_lista(sin, l*3) == aplicar_na_lista(sin, l)*3
```

## Aplicando várias vezes

Escreva uma função recursiva `aplicar_n_vezes` que recebe
- uma função `f`,
- uma lista `l`,
- e um inteiro `n`;

e retorna a lista resultante ao aplicar $n$ vezes a função `f` em cada elemento de `l`.

<!-- grade_id: aplicar_n_vezes, points: 0, grade: false, solution: true -->
```python
def aplicar_n_vezes(f,l,n=2):
    # YOUR CODE HERE
    raise NotImplementedError()
```

Testes básicos

<!-- grade_id: aplicar_n_vezes_2, points: 2, grade: true, solution: false -->
```python
assert aplicar_na_lista(odd, aplicar_na_lista(odd, l)) == aplicar_n_vezes(odd, l, 2)
```

<!-- grade_id: aplicar_n_vezes_3, points: 1, grade: true, solution: false -->
```python
assert aplicar_na_lista(sin, aplicar_na_lista(sin, aplicar_na_lista(sin, l))) == aplicar_n_vezes(sin, l, 3)
```

De novo, uma propriedade "geral":

<!-- grade_id: aplicar_n_vezes_35, points: 1, grade: true, solution: false -->
```python
assert aplicar_n_vezes(sin, aplicar_n_vezes(sin, l, 5), 3) == aplicar_n_vezes(sin, l, 8)
```

Vamos ver que você fez uma função recursiva, testando se a função chama a si mesma:

<!-- grade_id: aplicar_n_vezes_rec_call, points: 1, grade: true, solution: false -->
```python
assert has_selfcall(aplicar_n_vezes)
```

Também podemos testar que **de fato** a função está chamando a si mesma obrigando o Python a calcular "recursões demais",
o que gera um erro:

<!-- grade_id: aplicar_n_vezes_rec_explode, points: 2, grade: true, solution: false -->
```python
import sys
prevlim = sys.getrecursionlimit()
sys.setrecursionlimit(50)

try:
    aplicar_n_vezes(sin, l, 50)
    sys.setrecursionlimit(prevlim)
    assert False
except RuntimeError as e:
    sys.setrecursionlimit(prevlim)
    assert e.args[0][:32] == 'maximum recursion depth exceeded'
```

Agora, escreva uma versão _iterativa_ da função `aplicar_n_vezes`, que usa um `for`,
e não estoura o limite de recursão (já que a função não chama a si mesma!)

<!-- grade_id: aplicar_n_iter, points: 0, grade: false, solution: true -->
```python
def aplicar_n_vezes_iter(f,l,n=2):
    # YOUR CODE HERE
    raise NotImplementedError()
```

Agora, os testes são óbvios temos uma função "de referência" para comparar!

<!-- grade_id: aplicar_n_iter_1, points: 2, grade: true, solution: false -->
```python
assert aplicar_n_vezes_iter(sin, l, 3) == aplicar_n_vezes(sin, l, 3)
```

<!-- grade_id: aplicar_n_iter_2, points: 2, grade: true, solution: false -->
```python
res1 = aplicar_n_vezes_iter(cos, l2+l+l2, 8)
res2 = aplicar_n_vezes(cos, l2, 8) + aplicar_n_vezes(cos, l, 8) + aplicar_n_vezes(cos, l2, 8)
assert res1 == res2
```

Vejamos que a sua função não estoura a pilha!

<!-- grade_id: aplicar_n_iter_nonrec, points: 3, grade: true, solution: false -->
```python
sys.setrecursionlimit(50)

try:
    aplicar_n_vezes_iter(sin, l, 50)
    sys.setrecursionlimit(prevlim)
except RuntimeError as e:
    sys.setrecursionlimit(prevlim)
    assert False
```