![CC-BY-SA](https://mirrors.creativecommons.org/presskit/buttons/88x31/svg/by-sa.svg)
This notebook was created by [Bernardo Freitas Paulo da Costa](http://www.im.ufrj.br/bernardofpc),
and is licensed under Creative Commons BY-SA

# Teste para treinar
Este teste servirá como "teste" da plataforma de testes e provas para o curso de Computação Científica I. Assim, você poderá ter uma ideia de como serão corrigidas
- as listas de exercícios
- as provas

## Questão 1
Escreva uma função `squares` tal que, dado um inteiro $n \ge 0$, retorna a lista de todos os quadrados de $0$ até $n^2$.

<!-- grade_id: squares_function, points: 3, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
# Teste aqui a sua função antes de passar para a parte dos "asserts" abaixo.
squares(5) # Deveria produzir a lista [0, 1, 4, 9, 16, 25]
```

A caixa a seguir contém testes que serão efetuados sobre a sua função.
- Se algum teste falhar, você verá uma mensagem com `Assertion Error`.
- Senão, você verá a mensagem "Tudo ok". (por causa do "print"; em geral, será apenas vazio como nas seguintes)

Não apague: é ela que conta os pontos da questão!

<!-- grade_id: squares_test, points: 3, grade: true, solution: false -->
```python
assert(squares(0) == [0])
assert(squares(5) == [0, 1, 4, 9, 16, 25])
assert(squares(10) == [0, 1, 4, 9, 16, 25, 36, 49, 64, 81, 100])

print("Tudo ok")
```

## Questão 2
Escreva uma função `sqsum(n)` que retorna a soma de todos os quadrados de $1$ até $n$.

<!-- grade_id: sqsum_function, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
# Teste aqui a sua função.
# Você pode criar várias caixas de teste a mais sem problemas.
sqsum(10)
```

```python
sqsum(1000)
```

Novamente, a caixa abaixo contém testes para a sua função.

<!-- grade_id: sqsum_test, points: 3, grade: true, solution: false -->
```python
assert(sqsum(5) == 55)
assert(sqsum(10) == 385)
assert(sqsum(123) == 627874)
```

A caixa a seguir também contém testes, mas agora eu não digo quais serão.

Vale lembrar, mais uma vez, a mesma observação de antes: mantenha esta caixa: ela contém os pontos referentes aos testes, mesmo que você não os possa ver.

```python

```

<!-- grade_id: sqsum_test2, points: 3, grade: true, solution: false -->
```python

```

### Observação

Para ter certeza que este notebook será corrigido corretamente, o melhor é rodar tudo "de uma vez só, do zero".
Para simular isto, use o menu: `Kernel -> Restart & Run All`.