![CC-BY-SA](https://mirrors.creativecommons.org/presskit/buttons/88x31/svg/by-sa.svg)
This notebook was created by [Bernardo Freitas Paulo da Costa](http://www.im.ufrj.br/bernardofpc),
and is licensed under Creative Commons BY-SA

```python
import numpy as np
import matplotlib.pyplot as plt
```

<!-- grade_id: fourier_model, points: 0, grade: false, solution: true -->
```python
def model(cks, x):
    """Calcula f(x) para os coeficientes de Fourier c_k dados na lista cks.
    A lista cks tem comprimento ímpar, e os coeficientes vem na ordem [-M,...,-1,0,1,2,...,M]."""
    n = len(cks)
    assert(n%2 == 1)
    M = (n-1)//2

    # Atenção para que esta função seja vetorizada em x
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: fourier, points: 0, grade: false, solution: true -->
```python
def fourier(y):
    """Transformada discreta de Fourier, o vetor `y` corresponde aos valores em (2n+1) pontos igualmente espaçados.
    Retorna tanto o vetor dos coeficientes como uma função que calcula a interpolação."""
    n = len(y)
    assert(n%2 == 1)
    M = (n-1)//2
    # Calcule os coeficientes cks
    cks = np.zeros(n, dtype='complex')
    # YOUR CODE HERE
    raise NotImplementedError()
    return cks, (lambda x: model(cks,x))
```

# Interpolação via Fourier

Suponha que temos uma função $f$ num intervalo $[a,b]$,
e que calculamos $f(x)$ para $n=2M+1$ pontos $x_i$, igualmente espaçados de $a$ até $b$, inclusive.
(pois $f(b)$ pode ser diferente de $f(a)$, em geral!)

Podemos usar os coeficientes de Fourier para calcular uma função que aproxima $f$ neste intervalo.
- Qual o "período" de $f$, implícito nesta escolha de pontos?
- Como transformar o vetor $t$ para obtermos uma função $2\pi$ periódica?
- É preciso transformar mais alguma coisa?

<!-- grade_id: propriedades1, points: 3, grade: true, solution: true -->
YOUR ANSWER HERE

<!-- grade_id: fourier_interp, points: 0, grade: false, solution: true -->
```python
def fourier_interp(a,b,ys):
    """Interpolação de Fourier, com período no intervalo [a,b+h], dados 2m+1 valores igualmente espaçados entre a e b."""
    n = len(ys)
    assert(n%2 == 1)
    cks, fi = fourier(ys)
    # YOUR CODE HERE
    raise NotImplementedError()
```

### Asserts

<!-- grade_id: equal_nodes, points: 2, grade: true, solution: false -->
```python
xs = np.linspace(-1,2,num=21)
ys = np.exp(xs)
exp_interp = fourier_interp(-1,2,ys)
assert np.allclose(exp_interp(xs), ys, atol=1e-12, rtol=1e-12)
```

```python
ts = np.random.rand(10)
exp_interp(ts) - np.exp(ts)
```

<!-- grade_id: close_middle, points: 3, grade: true, solution: false -->
```python
xs = np.linspace(-1,2,num=21)
ys = np.exp(xs)
exp_interp = fourier_interp(-1,2,ys)

np.random.seed(1)
ts = np.random.rand(10)
assert np.allclose(exp_interp(ts), np.exp(ts), atol=0.2)
```

### Gráfico

Faça o gráfico da interpolação e do erro, na mesma figura.

<!-- grade_id: graf_interp_err, points: 4, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

### Oscilação

Porque você acha que houve as oscilações nos gráficos acima?

<!-- grade_id: explain_osc, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

## Comparando com Lagrange

```python
from interpolation import lagrange_v as lagrange
```

Faça o gráfico do erro de interpolação, nas mesmas condições, para a interpolação de Lagrange nestes 21 pontos.

<!-- grade_id: graf_err, points: 3, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

### Comente

<!-- grade_id: lagr_vs_fourier_1, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

# Refletindo um pouco

Em vez de interpolar os valores de $f(x)$ em $[a,b]$, vamos repetir os valores em ordem oposta.

### Intervalo efetivo

Qual era o período implícito de $f$ no caso anterior?
Qual o período agora?

<!-- grade_id: periodo_refletido, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

### Usando `np.fft`

Como vamos ter um número par de pontos, devido à repetição,
o mais simples será usar `np.fft.fft` para calcular os coeficientes $c_m$.

Dê a nova transformação afim a ser efetuada em $x$, dado o "novo período" de $f$,
e em seguida implemente a interpolação simétrica.

<!-- grade_id: t_afim_coefs, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

<!-- grade_id: fourier_interp_sym, points: 2, grade: true, solution: true -->
```python
def fourier_interp_sym(a,b,ys):
    n = len(ys)
    ysy = np.concatenate([ys, ys[::-1]])
    cks = np.fft.fft(ysy)
    # YOUR CODE HERE
    raise NotImplementedError()
```

Refaça os gráficos da interpolação e do erro:

<!-- grade_id: novos_grafs, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Comente.

<!-- grade_id: erros_interp_fourier_sym, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

# Mais uma interpolação

Compare a interpolação de Fourier (sem simetria) com a interpolação de Lagrange para
$$f(x) = \frac{1}{1 + x^2}$$
no intervalo $[-5,5]$.

<!-- grade_id: nova_f, points: 2, grade: true, solution: true -->
```python
# Defina $f$ e os pontos de interpolação
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: nova_f_fourier, points: 1, grade: true, solution: true -->
```python
# Faça os gráficos para Fourier
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: nova_f_lagr, points: 2, grade: true, solution: true -->
```python
# Enfim, os gráficos para Lagrange
# YOUR CODE HERE
raise NotImplementedError()
```

Comente

<!-- grade_id: final_comms, points: 3, grade: true, solution: true -->
YOUR ANSWER HERE