![CC-BY-SA](https://mirrors.creativecommons.org/presskit/buttons/88x31/svg/by-sa.svg)
This notebook was created by [Bernardo Freitas Paulo da Costa](http://www.im.ufrj.br/bernardofpc),
and is licensed under Creative Commons BY-SA

```python
import numpy as np
import matplotlib.pyplot as plt
```

# Calculando derivadas de ordem superior

A primeira ideia para calcular a derivada segunda de uma função é calcular duas vezes a derivada.
Inclusive porque esta é a definição!

Assim, $f''(x) \sim [f'(x+h) - f'(x-h)]/2h$, e o cálculo de $f'(x\pm h)$ pode ser feito analogamente.

Implemente esta fórmula.

<!-- grade_id: ddf, points: 0, grade: false, solution: true -->
```python
def ddf(f, x, h=1e-5):
    """Calcula uma aproximação da segunda derivada de $f$, nos pontos do vetor x, usando passos de tamanho $h$."""
    # YOUR CODE HERE
    raise NotImplementedError()
```

Testando uma função fácil.

<!-- grade_id: ddf1, points: 1, grade: true, solution: false -->
```python
# Perto
assert np.abs(ddf(np.sin, 1, 0.01) - (-np.sin(1))) < 1e-4
```

<!-- grade_id: ddf2, points: 1, grade: true, solution: false -->
```python
# Mas não perto demais
assert np.abs(ddf(np.sin, 1, 0.01) - (-np.sin(1))) > 1e-6
```

<!-- grade_id: ddf3, points: 1, grade: true, solution: false -->
```python
# Bem mais perto de algo um pouco errado
assert np.isclose(ddf(np.sin, 1, 0.01), -0.841442936, atol=1e-7)
```

Testando que, ao diminuir o passo, de fato o erro diminui.

<!-- grade_id: ddf4, points: 1, grade: true, solution: false -->
```python
assert np.abs(ddf(np.sin, 1, 0.001) - (-np.sin(1))) < 1e-6
```

<!-- grade_id: ddf5, points: 1, grade: true, solution: false -->
```python
assert np.abs(ddf(np.sin, 1, 0.001) - (-np.sin(1))) > 1e-8
```

<!-- grade_id: ddf6, points: 1, grade: true, solution: false -->
```python
assert np.isclose(ddf(np.sin, 1, 0.001), -0.8414704338, atol=1e-9)
```

### Derivadas laterais

Poderíamos ter calculado duas vezes as derivadas laterais.
Implemente esta solução também.

<!-- grade_id: ddf_l, points: 0, grade: false, solution: true -->
```python
def ddf_lateral(f, x, h=1e-5):
    """Calcula uma aproximação da segunda derivada de $f$, nos pontos do vetor x, usando passos de tamanho $h$."""
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: ddf_l1, points: 1, grade: true, solution: false -->
```python
assert np.abs(ddf_lateral(np.sin, 1, 0.01) + np.sin(1)) < 1e-2
assert np.abs(ddf_lateral(np.sin, 1, 0.01) + np.sin(1)) > 1e-3
```

<!-- grade_id: ddf_l2, points: 1, grade: true, solution: false -->
```python
assert np.isclose(ddf_lateral(np.sin, 1, 0.01), -0.84682478770914393, atol=1e-5)
```

Ao diminuir o passo, o erro diminui também

<!-- grade_id: ddf_l3, points: 2, grade: true, solution: false -->
```python
e1 = ddf_lateral(np.sin, 1, 0.01) + np.sin(1)
e2 = ddf_lateral(np.sin, 1, 0.001) + np.sin(1)
assert np.isclose(np.abs(e1/e2), 10, atol=1)
```

<!-- grade_id: ddf_l4, points: 2, grade: true, solution: false -->
```python
np.random.seed(1)
for x in np.random.randn(100):
    e1 = ddf_lateral(np.sin, x, 0.01) + np.sin(x)
    e2 = ddf_lateral(np.sin, x, 0.0001) + np.sin(x)
    assert np.isclose(np.abs(e1/e2), 100, atol=15), (x, e1, e2, e1/e2)
```

## Gráficos

Faça um gráfico do erro da segunda derivada, para algum valor do passo,
usando ambos os métodos `ddf` e `ddf_lateral`.
Qual dá o maior erro?

<!-- grade_id: graf_err, points: 3, grade: true, solution: true -->
```python
f = np.sin
xs = np.arange(0,7,0.05)
# YOUR CODE HERE
raise NotImplementedError()
```

Mude o tamanho do passo, e refaça os gráficos.
**Antes** de aparecer erro numérico, os erros tendem a zero?
Com que velocidade?

<!-- grade_id: graph_err2, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Comente abaixo a velocidade de decaimento do erro.

<!-- grade_id: err_decay, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

## Limites de passo

Já sabemos que não podemos fazer $h \to 0$.
O objetivo desta parte é encontrar um $h$ "quase ótimo" para cada um dos métodos.

```python
### Interactive solution
from ipywidgets import IntSlider, interactive
```

Escreva uma função `do(logh1, logh2)` que faz dois gráficos de erro do seno, no intervalo $[0,7]$,
para passos $h_1 = 2^{\tt logh1}$ para a derivada simétrica `ddf`,
e respectivamente para $h_2 = 2^{\tt logh2}$ para a derivada lateral `ddf_lateral`.

<!-- grade_id: do_iter, points: 3, grade: true, solution: true -->
```python
xs = np.arange(0,7,0.05)
# YOUR CODE HERE
raise NotImplementedError()
```

```python
logh1s = IntSlider(min=0, value=5, max=30)
logh2s = IntSlider(min=0, value=5, max=30)

interactive(do, logh1=logh1s, logh2=logh2s)
```

## Generalizando

Como seria a fórmula para a derivada terceira?
Implemente o método simétrico abaixo.

<!-- grade_id: dddf, points: 0, grade: false, solution: true -->
```python
def dddf(f, x, h=1e-5):
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: dddf1, points: 1, grade: true, solution: false -->
```python
assert (dddf(np.sin, 1, 0.01) + np.cos(1)) < 1e-4
```

<!-- grade_id: dddf2, points: 1, grade: true, solution: false -->
```python
assert (dddf(np.sin, 1, 0.01) + np.cos(1)) > 1e-5
```

Fixando $x=1$, e variando $h$, faça um gráfico do erro de cálculo para a derivada terceira do seno.
Qual parece ser o melhor passo para a derivada terceira simétrica?

<!-- grade_id: graph_err_ddd1, points: 2, grade: true, solution: true -->
```python
# Seno
# YOUR CODE HERE
raise NotImplementedError()
```

Faça abaixo um gráfico análogo, para a derivada terceira da exponencial, calculada em $x = 1, 3, 10$.
Use uma única figura e três curvas, portanto não esqueça da legenda.

<!-- grade_id: graph_err_ddd2, points: 3, grade: true, solution: true -->
```python
xs = np.array([0.1,0.3,1,3])
# YOUR CODE HERE
raise NotImplementedError()
```

Comente abaixo os gráficos em especial o ponto de menor erro para a derivada terceira.

Você percebe alguma relação do "melhor $h$" para cada uma das derivadas simétricas que calculamos ($f'$, $f''$ e $f'''$)?

<!-- grade_id: comments, points: 3, grade: true, solution: true -->
YOUR ANSWER HERE