# Estimativas empíricas de erros

Gostaríamos de obter estimativas do erro obtido com a solução numérica de uma EDO.
Devido à natureza cumulativa dos erros no método de Euler, é razoável usar o ponto final,
logo uma primeira idéia é comparar o valor no ponto final entre duas soluções,
sendo uma com um número maior de pontos que a outra.

A motivação deste método é que, se $y_1$ e $y_2$ são as estimativas do ponto final real $y$,
com $|y_1 - y| < \varepsilon_1$ e $|y_2 - y| < \varepsilon_2$, então $|y_1 - y_2| < \varepsilon_1 + \varepsilon_2$.
Entretanto, temos apenas como calcular $|y_1 - y_2|$,
o que não garante (matematicamente) nada quanto aos valores que desejamos realmente calcular,
$|y_1 - y|$ ou $|y_2 - y|$, mas pode dar uma indicação de sua ordem de grandeza.

Veremos ao longo deste teste algumas das limitações deste método.

## Bibliotecas

Além do numpy e da matplotlib, vamos precisar da `euler_npts` de uma aula anterior.
Copie aqui o código da mesma (e da `eulerexplicito`, necessária para ela).

```python
import numpy as np
import matplotlib.pyplot as plt
```

<!-- grade_id: cell-59ed968e77d270f9, points: 0, grade: false, solution: true -->
```python
def eulerexplicito(F, t0, y0, ts):
    """Calcula uma solução aproximada da equação y' = F(t,y) pelo método de Euler, nos pontos [ts]."""
    # YOUR CODE HERE
    raise NotImplementedError()

def euler_npts(F, I, y0, npts, retpts=False):
    """Calcula uma solução aproximada da equação  y' = F(t,y)  pelo método de Euler,
    com `npts` pontos igualmente espaçados no intervalo `I`."""
    # YOUR CODE HERE
    raise NotImplementedError()
```

## Questão 1: Estimando o erro empiricamente.

Implemente a função `estima_erro(F, I, y0, n1, n2)` que retorna a estimativa do erro a partir das soluções da equação diferencial $y'(t) = F(t, y(t))$ com `n1` pontos e `n2` pontos igualmente espaçados.

<!-- grade_id: cell-167d9d30f0958df0, points: 0, grade: false, solution: true -->
```python
def estima_erro(F, I, y0, n1, n2):
    # YOUR CODE HERE
    raise NotImplementedError()
```

### Alguns testes simples

<!-- grade_id: cell-0466da9ea29d4cee, points: 1, grade: true, solution: false -->
```python
def f_exp(t,x): return x
a1 = estima_erro(f_exp, [0,1], 1.2, 10,100)
assert np.isclose(a1, 0.148, atol=0.01)
```

<!-- grade_id: cell-3756d643be1f5178, points: 1, grade: true, solution: false -->
```python
def f_mexp(t,x): return -x
a2 = estima_erro(f_mexp, [0,1], 2.2, 10,100)
assert np.isclose(a2, 0.043, atol=0.01)
```

<!-- grade_id: cell-5af53e64f1c2f3e5, points: 1, grade: true, solution: false -->
```python
def f_poly(t,x): return t**2 - 2*t
a3 = estima_erro(f_poly, [-4,4], 100, 200,2000)
assert np.isclose(a3, 0.292, atol=0.01)
```

### E um gráfico

Faça o gráfico da estimativa do erro para 10 e $n$ pontos, variando $n$ de 10 a 50.

Dê um título para seu gráfico!

<!-- grade_id: cell-3e26b5a6494fab3b, points: 2, grade: true, solution: true -->
```python
ns = np.arange(10,50)
# YOUR CODE HERE
raise NotImplementedError()
plt.show()
```

### Análise

O erro parece crescer ou diminuir conforme $n$ aumenta? Comente/explique.

<!-- grade_id: cell-3e3f15e97941844a, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

## Questão 2:

Agora que temos uma estimativa de erro podemos construir um integrador que retorna a solução dentro de uma tolerância especificada a priori.

Implemente a função `dobrador()` que calcula duas soluções com `n` e `2n` pontos respectivamente e caso o erro estimado esteja abaixo da tolerância `tol` retorna o par `(ts, ys)` da solução com mais pontos, caso contrário ele tenta o mesmo procedimento dobrando o valor inicial de `n`.

**Dica:** Você pode modificar a função `estima_erro` para receber apenas $n$, e retornar os tempos de integração e a lista de pontos da solução com mais pontos.

<!-- grade_id: cell-747b1c8864df66b7, points: 0, grade: false, solution: true -->
```python
def estima_erro_novo(F, I, y0, n):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-5b38d71b216fc96f, points: 0, grade: false, solution: true -->
```python
def dobrador(F, I, y0, n, tol, numit=False):
    """Calcula uma solução aproximada da equação y'(t) = F(t,y(t)) no intervalo I, com condição inicial y0,
    dobrando o número de pontos intermediários até que a estimativa de erro seja < `tol`.
    
    Se `numit` = True, retorna o número de vezes que dobramos o número de pontos intermediários."""
    num_it = 0
    # YOUR CODE HERE
    raise NotImplementedError()
```

Testes simples do dobrador: como temos apenas uma *estimativa* do erro, o teste de que está perto da solução real tem uma margem de segurança ;-)

<!-- grade_id: cell-c82c76c5f9e2e968, points: 2, grade: true, solution: false -->
```python
ts, ys = dobrador(f_exp, [0,2], 2, 1, 1e-3)
assert ys[0] == 2
assert np.isclose(ys[-1], 2*np.exp(2), atol=2e-3)
```

<!-- grade_id: cell-5f171e034d33f018, points: 1, grade: true, solution: false -->
```python
ts, ys = dobrador(f_mexp, [0,2], 2, 1, 1e-4)
assert ys[0] == 2
assert np.isclose(ys[-1], 2*np.exp(-2), atol=2e-4)
```

<!-- grade_id: cell-c83ad88d7718d4ee, points: 1, grade: true, solution: false -->
```python
ts1, ys1, n1 = dobrador(f_mexp, [0,2], 2, 1, 1e-4, numit=True)
ts2, ys2, n2 = dobrador(f_mexp, [0,2], 2, 1, 2e-4, numit=True)

assert n1 - n2 == 1
assert len(ts1) == 2*len(ts2)
```

## Questão 3a: Exponencial

Vamos verificar que esta ideia funciona para a EDO da família exponencial, $y'(t) = y(t)$.

Faça o gráfico da solução numérica no intervalo $[0,2]$, com condição inicial $y(0) = 3/4$,
tentando obter uma precisão de 2 casas decimais, e desenhe também a solução exata no mesmo gráfico.
Lembre de botar:
- legenda
- título

<!-- grade_id: cell-d3e0e03acbb107d5, points: 0, grade: false, solution: true -->
```python
# Abaixo, faça as contas (com o `dobrador`) e dê os comandos de gráfico (plt.plot, etc).
# A caixa seguinte contém alguns asserts, então é importante que você não use `plt.show()` no seu código,
# pois senão o objeto `ax` estaria vazio...

# YOUR CODE HERE
raise NotImplementedError()

ax = plt.gca()
plt.show()
```

<!-- grade_id: cell-ed62e72005956f1a, points: 1, grade: true, solution: false -->
```python
assert ax.title.get_text() != ""
assert len(ax.lines) == 2
assert len(ax.legend().texts) == 2
```

<!-- grade_id: cell-d927a3873fe3a154, points: 2, grade: true, solution: false -->
```python
# Testa que as curvas estão próximas da solução
tf = []
for l in ax.lines:
    lx, ly = l.get_xdata(), l.get_ydata()
    assert np.allclose(ly, 3/4*np.exp(lx), atol=2e-2)
    tf.append( np.allclose(ly, 3/4*np.exp(lx), atol=1e-8) )

# E que só UMA delas está próxima, a outra está um pouco longe porque é numérica
assert all(tf) == False
assert any(tf) == True
```

Agora, faça o gráfico do **erro**: não esqueça do título.

<!-- grade_id: cell-d3e8b20c31c68dac, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

ax = plt.gca()
plt.show()
```

<!-- grade_id: cell-227abac7a38b508c, points: 2, grade: true, solution: false -->
```python
assert ax.title.get_text() != ""
assert len(ax.lines) == 1
ydata = ax.lines[0].get_ydata()
assert min(ydata) > -0.01
assert max(ydata) <  0.01
```

Quantos pontos você precisou usar neste caso?
Dê sua resposta na caixa abaixo, sob a forma de código python que calcule este valor a partir dos valores retornados pela função `dobrador()`.

<!-- grade_id: cell-c05340e660325b73, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Questão 3b: Exponencial decrescente

Vejamos o que acontece se agora tentamos resolver a equação diferencial $y'(t) = - y(t)$, cuja solução é $e^{-t}y_0$.
Faça os gráficos tanto das soluções como do erro da solução, no mesmo intervalo $[0,2]$ e para a mesma condição inicial,
também almejando uma precisão de 2 casas decimais.

<!-- grade_id: cell-57e3c26c7cc52bc4, points: 0, grade: false, solution: true -->
```python
# Faça contas aqui
# YOUR CODE HERE
raise NotImplementedError()

ax = plt.gca()
plt.show()
```

<!-- grade_id: cell-06406184fd88c4c6, points: 1, grade: true, solution: false -->
```python
assert ax.title.get_text() != ""
assert len(ax.lines) == 2
assert len(ax.legend().texts) == 2
```

<!-- grade_id: cell-c5c3c7ed0b89121f, points: 1, grade: true, solution: false -->
```python
# Testa que as curvas estão próximas da solução
tf = []
for l in ax.lines:
    lx, ly = l.get_xdata(), l.get_ydata()
    assert np.allclose(ly, 3/4*np.exp(-lx), atol=2e-2)
    tf.append( np.allclose(ly, 3/4*np.exp(-lx), atol=1e-8) )

# E que só UMA delas está próxima, a outra está um pouco longe porque é numérica
assert all(tf) == False
assert any(tf) == True
```

E o erro:

<!-- grade_id: cell-39ccedaabd676921, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

ax = plt.gca()
plt.show()
```

<!-- grade_id: cell-d19bbb4fce578c3b, points: 1, grade: true, solution: false -->
```python
assert ax.title.get_text() != ""
assert len(ax.lines) == 1
ydata = ax.lines[0].get_ydata()
assert min(ydata) > -0.01
assert max(ydata) <  0.01
```

Quantos pontos foram usados agora?

<!-- grade_id: cell-4c2868f4cb2b7a6d, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

### Análise

Comente a diferença que você observou, tanto no número de pontos, quanto no comportamento do erro.

<!-- grade_id: cell-ea556388b8bfcad8, points: 3, grade: true, solution: true -->
YOUR ANSWER HERE

## Questão 4: Um defeito da estimativa

Considere as "equações diferenciais" $y'(t) = \cos(t)$ e $z'(t) = \sin(t)$, cujas soluções são, respectivamente,
$y(t) = y(0) + \sin(t)$ e $z(t) = z(0) - \sin(t)$.

Observe atentamente os `assert`s abaixo:

```python
def F1(t,x):
    return np.cos(t)

def F2(t,x):
    return np.sin(t)

assert estima_erro(F1, [     0,2*np.pi], 0, n1=3, n2=5) < 1e-15
assert estima_erro(F2, [-np.pi,  np.pi], 0, n1=3, n2=5) < 1e-15
```

```python
for k1, k2 in zip([-2,-1,0,1],[-1,0,1,2]):
    assert estima_erro(F1, [2*k1*np.pi,2*k2*np.pi], 0, n1=3, n2=5) < 1e-15
    assert estima_erro(F2, [2*k1*np.pi,2*k2*np.pi], 0, n1=3, n2=5) < 1e-15
```

Aparentemente, usar 3 ou 5 pontos dá praticamente o mesmo valor para a solução ao final de um período da função.

Faça o gráfico da solução aproximada da primeira equação para estes dois valores de $n$ no intervalo $[0,2\pi]$,
e inclua também a solução $\sin(t)$.
Dê uma legenda e um título!

<!-- grade_id: cell-0852ce747c7f9dcb, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

ax = plt.gca()
plt.show()
```

<!-- grade_id: cell-ab5b172ad5eda771, points: 1, grade: true, solution: false -->
```python
assert ax.title.get_text() != ""
assert len(ax.lines) == 3
assert len(ax.legend().texts) == 3
```

<!-- grade_id: cell-b736e67da79ad345, points: 2, grade: true, solution: false -->
```python
npts_sorted = sorted([len(l.get_xdata()) for l in ax.lines])
assert npts_sorted[0] == 3
assert npts_sorted[1] == 5
assert npts_sorted[2] > 20
```

Agora, faça o gráfico dos erros:

<!-- grade_id: cell-6b2e64caac7446ce, points: 0, grade: false, solution: true -->
```python
I = [0,2*np.pi]
y0 = 0

# YOUR CODE HERE
raise NotImplementedError()

ax = plt.gca()
plt.show()
```

<!-- grade_id: cell-7ddab8fb5c8a00e5, points: 1, grade: true, solution: false -->
```python
assert ax.title.get_text() != ""
assert len(ax.lines) == 2
assert len(ax.legend().texts) == 2
```

Faça o gráfico da solução aproximada (e do erro, se quiser) para ("muito") mais pontos.

<!-- grade_id: cell-7bfb3e3f30e7bbb4, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
plt.show()
```

### Análise

Porque a estimativa de erro calculada pelo valor final é tão pequena?
Como seria possível contornar este problema, para termos uma noção melhor de que a solução já está "boa o suficiente"?

<!-- grade_id: cell-256b52f820f651b4, points: 3, grade: true, solution: true -->
YOUR ANSWER HERE