![CC-BY-SA](https://mirrors.creativecommons.org/presskit/buttons/88x31/svg/by-sa.svg)
This notebook was created by [Bernardo Freitas Paulo da Costa](http://www.im.ufrj.br/bernardofpc),
and is licensed under Creative Commons BY-SA

```python
import numpy as np
import matplotlib.pyplot as plt
```

# Calculando primitivas

Agora que sabemos calcular integrais, também podemos calcular primitivas de funções.
Há algumas ideias de como fazer isso:
- uma função `primitiva(f,x,h)`, análoga a `df(f,x,h)`, que calcula a primitiva de $f$ nos pontos do vetor `x`.
- uma função `primitiva(f,h)`, análoga a `df(f,h)`, que retorna uma função que calcula a primitiva de $f$ "sob demanda".

Mas isso não basta:
Imagine, por exemplo, que vamos calcular a primitiva de $\sin(x^2)$, chame esta função de $F$.
Ora, o valor de $F(1)$ não é único, ele depende de uma _constante de integração_,
que pode ser assimilada ao ponto inicial da integral (numérica!) que vamos fazer.
Assim, nossos métodos precisam de um argumento a mais: $a$, o ponto inicial.

Temos mais uma dificuldade a vencer: a estrutura da nossa função supõe que fixamos o passo $h$ de integração.
O que acontece se o usuário quiser calcular a primitiva em $a + 23.512h$???

Uma solução diferente seria fixar o número de subdivisões, $n$.
Mas isso também não é muito bom:
Se $n$ for grande, mas estivermos calculando a primitiva num ponto próximo, vamos estar "fazendo contas à toa";
se $n$ não for grande o suficiente, ao calcular num ponto mais distante, a resposta estará muito errada.

## Organização

- A primeira parte da prova irá construir as aproximações de $F$, calculando as primitivas "com passo $h$".
- A segunda parte irá desenvolver técnicas de interpolação para resolver o problema de $a+ 23.512h$.

Por questões de tempo (e tamanho da prova!) a terceira parte, que traça as primitivas para valores fracionários de $h$,
e a quarta parte, que se esforça para calcular $f$ o mínimo possível, ficam para uma outra vez ;-)

# Parte 1: Primitivas (35 pts)

## Ferramentas básicas

<!-- grade_id: int_meth, points: 0, grade: false, solution: true -->
```python
# Inclua aqui cauchy, midpoint e simpson
# YOUR CODE HERE
raise NotImplementedError()
```

```python
ans = 1 - np.cos(1)
a =   cauchy(np.sin, 0, 1, n=1000)
b = midpoint(np.sin, 0, 1, n=500)
c =  simpson(np.sin, 0, 1, n=100)
np.array([a, b, c]) - ans
```

## 1.1 A função `primitiva()` simples

As funções de integração recebem $n$, então é preciso "converter" o argumento $h$
(que faz mais sentido para nós neste caso) em $n$.
Este é o propósito da função `primitiva()` abaixo.

Obs: não é óbvio como vetorizar a função `F()` abaixo.  Não vamos precisar disso.

<!-- grade_id: primitiva, points: 0, grade: false, solution: true -->
```python
def primitiva(f,a,h=0.01,meth=simpson):
    """Retorna uma função que calcula a integral de f no intervalo [a,x], usando o método meth."""
    def F(x):
    # YOUR CODE HERE
    raise NotImplementedError()
    return F
```

```python
def veryclose(a,b):
    return np.isclose(a,b, rtol=1e-12, atol=1e-12)
```

Testando o funcionamento básico de `primitiva`

<!-- grade_id: primitiva_t1, points: 1, grade: true, solution: false -->
```python
F = primitiva(np.sin,0)
assert veryclose(F(1), simpson(np.sin,0,1))
```

<!-- grade_id: primitiva_t2, points: 1, grade: true, solution: false -->
```python
F = primitiva(np.exp,0)
assert veryclose(F(1), simpson(np.exp,0,1))
```

<!-- grade_id: primitiva_t3, points: 1, grade: true, solution: false -->
```python
F = primitiva(np.sin,0)
assert veryclose(F(10), simpson(np.sin,0,10,n=1000))
```

Testando que os argumentos são levados em conta

<!-- grade_id: primitiva_t4, points: 1, grade: true, solution: false -->
```python
F = primitiva(np.sin,4)
assert veryclose(F(5), simpson(np.sin,4,5))
```

<!-- grade_id: primitiva_t5, points: 1, grade: true, solution: false -->
```python
F = primitiva(np.exp,0, meth=cauchy)
assert veryclose(F(1), cauchy(np.exp,0,1))
```

<!-- grade_id: primitiva_t6, points: 1, grade: true, solution: false -->
```python
F = primitiva(np.cos,0, h=0.1)
assert veryclose(F(1), simpson(np.cos,0,1,n=10))
```

Testes finais

<!-- grade_id: primitiva_t7, points: 1, grade: true, solution: false -->
```python
F = primitiva(np.sin,10, h=0.001, meth=midpoint)
assert veryclose(F(12), midpoint(np.sin,10,12,n=2000))
```

<!-- grade_id: simp_neg, points: 1, grade: true, solution: false -->
```python
assert simpson(np.cos, 1, 0) < 0
```

<!-- grade_id: primitiva_t8, points: 2, grade: true, solution: false -->
```python
F = primitiva(np.cos,2)
assert veryclose(F(1), simpson(np.cos,2,1))
```

## 1.2 Erros

Agora, vamos estudar como o erro da primitiva evolui,
- ao afastar o ponto final $x$ de $a$; e
- ao aumentar / diminuir $h$.

```python
F_num = primitiva(np.sin,0)
```

<!-- grade_id: F_analitica, points: 1, grade: true, solution: true -->
```python
# Dê aqui a primitiva exata correspondente à expressão acima:
def F_analitica(x):
    # YOUR CODE HERE
    raise NotImplementedError()
```

Faça o gráfico do erro de `F_num` para `F_analítica` com $x \in [5,8]$.

<!-- grade_id: graph_err, points: 2, grade: true, solution: true -->
```python
xs = np.linspace(5,8,num=100)
# YOUR CODE HERE
raise NotImplementedError()
```

Faça agora no intervalo $[5,100]$.
Tome cuidado para o gráfico ficar claro!

<!-- grade_id: graph_err_100, points: 3, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Comente

<!-- grade_id: comm_per, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

Mantendo o intervalo $[5,8]$,
trace as curvas de erro (no mesmo eixo) para 3 valores diferentes de $h$.

<!-- grade_id: graph_3err, points: 3, grade: true, solution: true -->
```python
xs = np.linspace(5,8,num=100)
# YOUR CODE HERE
raise NotImplementedError()
```

Faça também o gráfico para outro método de integração.

<!-- grade_id: graph_other_meth, points: 2, grade: true, solution: true -->
```python
xs = np.linspace(5,8,num=100)
# YOUR CODE HERE
raise NotImplementedError()
```

A taxa de decaimento é a que você espera? Porquê?

<!-- grade_id: decay_order, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

## 1.3 Mais erros

O que a sua função faz quando o tamanho do intervalo de integração não é múltiplo de $h$?

<!-- grade_id: graph_04, points: 1, grade: true, solution: true -->
```python
xs = np.linspace(5,8,num=100)
h = 0.4
F_num = primitiva(np.sin,0, h=h)
# YOUR CODE HERE
raise NotImplementedError()
```

Você consegue explicar este comportamento do erro?

<!-- grade_id: err_04_explain, points: 3, grade: true, solution: true -->
YOUR ANSWER HERE

Faça novos gráficos do erro (com valores diferentes para $h$, e com métodos de integração diferentes).
Isso apóia a sua explicação acima?

<!-- grade_id: graph_explain_04, points: 4, grade: true, solution: true -->
```python
xs = np.linspace(5,8,num=100)
_ , axs = plt.subplots(ncols=3, figsize=(15,4))

for m,ax in zip([cauchy,midpoint,simpson], axs):
    for h in [0.1, 0.3, 0.6]:
        # YOUR CODE HERE
        raise NotImplementedError()
    ax.set_title(m.__name__)
    ax.legend()
plt.show()
```

Comente abaixo.

<!-- grade_id: final_comm, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

# Parte 2: Interpolação parcial (25 pts)

Vimos que tomar uma interpolação de ordem muito alta pode ser muito difícil, e sujeito a erro.
Além disso, estamos quase sempre usando pontos igualmente espaçados para calcular a integral,
o que pode aumentar a sensibilidade ao erro.

Assim, vamos usar uma técnica "simples" de interpolação por partes
para contornar o problema.

## Ferramentas básicas

<!-- grade_id: lagrange, points: 0, grade: false, solution: true -->
```python
# interpolador de lagrange
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: lagrange_ok, points: 0, grade: true, solution: false -->
```python
# Isso vale 0 pontos, claro...
p = lagrange([0,1,2], [2,4,8])
assert p(3) == 14
```

## 2.1 Achando pontos próximos

Suponha que temos acesso a todos os valores de $y$ correspondentes a $x$ que são múltiplos de $0.1$.
Suponha também que vamos interpolar no intervalo $[0, 0.1]$.
Dê uma função que retorna os $n$ nós (xs) mais próximos de $0$.

<!-- grade_id: close_nodes, points: 0, grade: false, solution: true -->
```python
def close_nodes(n):
    """Retorna os n nós mais próximos de [0,0.1] com espaçamento 0.1"""
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
close_nodes(10)
```

```python
close_nodes(11)
```

<!-- grade_id: close_nodes_3, points: 1, grade: true, solution: false -->
```python
ans = np.array([-0.1,  0. ,  0.1])
assert np.allclose(close_nodes(3), ans, rtol=1e-12, atol=1e-12)
```

<!-- grade_id: close_nodes_11, points: 2, grade: true, solution: false -->
```python
ans = np.array([-0.5, -0.4, -0.3, -0.2, -0.1,  0. ,  0.1,  0.2,  0.3,  0.4,  0.5])
assert np.allclose(close_nodes(11), ans, rtol=1e-12, atol=1e-12)
```

<!-- grade_id: close_nodes_10, points: 2, grade: true, solution: false -->
```python
# Tanto faz se começa com -0.5 ou termina com 0.5
assert veryclose( np.sum(np.abs(close_nodes(10))), 2.5 )
```

## 2.2 Interpolador de pontos próximos

Agora, faça uma função `interp_prox(f,n)` que retorna o polinômio interpolador da função $f$
nos $n$ pontos mais próximos do intervalo $[0,0.1]$ como acima.

(Obs: na "vida real", teríamos uma lista de pontos `xs` e `ys`
dos quais vamos escolher os $n$ pares correspondentes aos $n$ $x_i$ mais próximos de cada $t$,
mas isto demanda mais programação, e não é tão importante assim para o que vamos fazer)

<!-- grade_id: interp_prox, points: 0, grade: false, solution: true -->
```python
def interp_prox(f,n):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
pol = interp_prox(np.sin, 5)
pol(0.01), pol(0.01) - np.sin(0.01)
```

<!-- grade_id: interp_prox_001, points: 2, grade: true, solution: false -->
```python
pol = interp_prox(np.sin, 5)
assert np.isclose(pol(0.01), np.sin(0.01))
assert not(veryclose(pol(0.01), np.sin(0.01)))
```

<!-- grade_id: interp_prox_004, points: 1, grade: true, solution: false -->
```python
pol = interp_prox(np.exp, 5)
assert np.isclose(pol(0.04), np.exp(0.04))
assert not(veryclose(pol(0.04), np.exp(0.04)))
```

<!-- grade_id: interp_prox_001_10, points: 2, grade: true, solution: false -->
```python
pol = interp_prox(np.sin, 10)
assert veryclose(pol(0.01), np.sin(0.01))
```

<!-- grade_id: interp_prox_007_10, points: 1, grade: true, solution: false -->
```python
pol = interp_prox(np.cos, 10)
assert veryclose(pol(0.07), np.cos(0.07))
```

## 2.3 Erros

Faça agora um gráfico do erro de interpolação da função `np.sin()` para $t$ no intervalo $[0,0.1]$,
para diferentes valores de $n$.

<!-- grade_id: graph_err_0_1, points: 3, grade: true, solution: true -->
```python
ns = [2,3,4,5,6]
# YOUR CODE HERE
raise NotImplementedError()
```

Repita para as funções exponencial e cosseno.
Faça dois gráficos na mesma figura.

<!-- grade_id: graph_err_exp_cos, points: 3, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

O que você percebe? Parece haver alguma explicação?

<!-- grade_id: compare_sin_cos_exp, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Como você faria para testar se sua explicação está correta?
Escolha uma outra função (ou _outras funções_!) para interpolar.

<!-- grade_id: interp_other, points: 3, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

E comente abaixo.

<!-- grade_id: interp_other_explain, points: 3, grade: true, solution: true -->
YOUR ANSWER HERE