![CC-BY-SA](https://mirrors.creativecommons.org/presskit/buttons/88x31/svg/by-sa.svg)
This notebook was created by [Bernardo Freitas Paulo da Costa](http://www.im.ufrj.br/bernardofpc),
and is licensed under Creative Commons BY-SA

```python
import numpy as np
import matplotlib.pyplot as plt
```

# Parte 1: Interpolação em domínios infinitos

Para calcular integrais em domínios infinitos, há duas técnicas:
- interpolar funções em $(-\infty,\infty)$ (ou $[0,+\infty)$, mas não vamos considerar este caso aqui);
- realizar uma mudança de variáveis para um intervalo finito, e interpolar em $(-1,1)$.

Nesta parte, vamos estudar a construção análoga à base de polinômios ortogonais para intervalos infinitos.
Como base de interpolação, vamos usar polinômios "corrigidos" para decairem no infinito:
Senão, eles não seriam integráveis.

Assim, vamos definir o produto interno entre dois polinômios $P$ e $Q$ por
$$ \langle P,Q\rangle = \int_{-\infty}^\infty P(t)Q(t) e^{-t^2/2}\, dt. $$

### 1.1: Integrais e "fatoriais"

Integrando por partes com $t e^{-t^2/2}$, mostre que
$$\int_{-\infty}^\infty t^{m+2} e^{-t^2/2} \, dt = (m+1) \int_{-\infty}^\infty t^m e^{-t^2/2} \, dt.$$

Como fica a fórmula acima em função de $n = m+2$?
(isso será útil para implementar o cálculo destas integrais, mais abaixo)

<!-- grade_id: ipp, points: 3, grade: true, solution: true -->
YOUR ANSWER HERE

É dado $\int_{-\infty}^\infty e^{-t^2/2} \, dt = C = \sqrt{2\pi}$.
Mostre que $\int_{-\infty}^\infty t e^{-t^2/2} \, dt = 0$.

<!-- grade_id: integracao, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

Agora, escreva a função `allints(n)`,
que calcula todas as integrais $\int_{-\infty}^\infty t^j e^{-t^2/2} \, dt$ para $0 \le j < n$.

Cuidado com a recorrência / _loop_:
a cada etapa diminuímos **2** do expoente.

<!-- grade_id: allints, points: 0, grade: false, solution: true -->
```python
C = np.sqrt(2*np.pi)
def allints(n):
    """Retorna uma lista de tamanho n, com as integrais de 0 até n-1."""
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
allints(5)
```

```python
allints(6)
```

<!-- grade_id: allints_6, points: 2, grade: true, solution: false -->
```python
assert allints(6) == [2.5066282746310002, 0, 2.5066282746310002, 0, 7.5198848238930012, 0]
```

<!-- grade_id: allints_rand, points: 2, grade: true, solution: false -->
```python
for n in np.random.randint(low=10,high=30,size=5):
    assert len(allints(int(n))) == n
```

```python
%time x = allints(30)
```

O teste abaixo deve passar em menos de 30 segundos.
Se a geração de 30 integrais (linha acima) já levar mais de 1 segundo, isso poderá ser um problema.

<!-- grade_id: allints_100, points: 2, grade: true, solution: false -->
```python
prev = 0
for i,x in enumerate(allints(100)):
    if i % 2 == 0:
        assert prev <= x
        prev = x
    else:
        assert x == 0    
```

### 1.2: Calculando o produto de polinômios

Escreva uma função que calcula os coeficientes do produto de dois polinômios.

<!-- grade_id: coefs_product, points: 0, grade: false, solution: true -->
```python
# De uma aula anterior
def coefs_product(p, q):
    """Calcula os coeficientes do polinômio P*Q.
    
    As listas / vetores de coeficientes p e q não precisam ter o mesmo tamanho."""
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
coefs_product([0,1,2],[0,0,1])
```

<!-- grade_id: coefprod_simple, points: 1, grade: true, solution: false -->
```python
ans = np.array([0, 0, 0, 1, 2])
assert np.all(ans == coefs_product([0,1,2],[0,0,1]))
```

```python
coefs_product([1,2,5],[3,4])
```

<!-- grade_id: coefprod_mix, points: 2, grade: true, solution: false -->
```python
ans = np.array([ 3, 10, 23, 20])
assert np.all(ans == coefs_product([1,2,5],[3,4]))
```

<!-- grade_id: coefprod_deg, points: 2, grade: true, solution: false -->
```python
for d1,d2 in np.random.randint(4,10, size=(10,2)):
    c1 = np.random.rand(d1+1) # Polinômio de grau d1
    c2 = np.random.rand(d2+1) # Polinômio de grau d2
    assert len(coefs_product(c1,c2)) == d1+d2+1 # Polinômio de grau d1+d2
```

### 1.3 Produto interno

Escreva a função `inner_expdecay` que calcula o produto interno de dois polinômios $P$ e $Q$,
dados pela lista de seus coeficientes.
Note que o cálculo das integrais dos monômios é dado pela função `allints()`!

<!-- grade_id: expdecay, points: 0, grade: false, solution: true -->
```python
def inner_expdecay(ps,qs):
    """Produto interno dado pela integral com decaimento exponencial."""
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: expdecay_1, points: 1, grade: true, solution: false -->
```python
assert inner_expdecay([0,0,1],[0,0,0,0,1]) == 15*np.sqrt(2*np.pi)
```

<!-- grade_id: expdecay_2, points: 1, grade: true, solution: false -->
```python
assert inner_expdecay([0,0,1],[0,1,0,1]) == 0
```

<!-- grade_id: expdecay_3, points: 1, grade: true, solution: false -->
```python
assert inner_expdecay([1,-1,0,1],[-1,0,1]) == 0
```

<!-- grade_id: expdecay_4, points: 1, grade: true, solution: false -->
```python
assert np.isclose(inner_expdecay([1,-1,0,1],[-1/4,0,0.3,0,1]), 3.05*C, rtol=1e-12, atol=1e-12)
```

### 1.4 Base de polinômios

```python
# Código para Gram-Scmidt
def resto(v, us, inner):
    prods = [inner(v,u)*u for u in us]
    vv = v - np.sum(prods, axis=0)
    return vv/np.sqrt(inner(vv,vv))
    
def gram_schmidt(vs, inner):
    """Retorna vetores ortognais com relação ao produto interno `inner`,
    a partir de uma coleção de vetores independentes em `vs`."""
    us = []
    while len(vs) > 0:
        v = vs[0]
        vv = resto(v, us, inner)
        us.append(vv)
        vs = vs[1:]
    return us
```

Escreva a função `findbase`, que retorna os $n$ primeiros polinômios ortogonais para o produto interno `inner`.

<!-- grade_id: findbase, points: 0, grade: false, solution: true -->
```python
def findbase(n,inner):
    """Retorna os coeficientes dos $n$ primeiros polinômios ortogonais para o produto interno dado por `inner`."""
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
base_3 = findbase(3, inner=inner_expdecay)
for deg, p in enumerate(base_3):
    print('{} -> {}'.format(deg, p))
```

<!-- grade_id: findbase_norm1, points: 2, grade: true, solution: false -->
```python
base_3 = findbase(3, inner=inner_expdecay)
for p in base_3:
    assert np.isclose(inner_expdecay(p,p), 1, rtol=1e-12, atol=1e-12)
```

<!-- grade_id: findbase_orth, points: 1, grade: true, solution: false -->
```python
base_3 = findbase(3, inner=inner_expdecay)
assert np.isclose(inner_expdecay(base_3[0], base_3[1]), 0, rtol=1e-12, atol=1e-12)
```

<!-- grade_id: findbase_orth2, points: 1, grade: true, solution: false -->
```python
base_3 = findbase(3, inner=inner_expdecay)
assert np.isclose(inner_expdecay(base_3[0], base_3[2]), 0, rtol=1e-12, atol=1e-12)
assert np.isclose(inner_expdecay(base_3[2], base_3[1]), 0, rtol=1e-12, atol=1e-12)
```

### 1.5 Gráficos

Faça o gráfico dos polinômios ortogonais de graus 7 e 8.
Faça, também, o gráfico das funções interpoladoras correspondentes,
que incorporam o fator $\exp(-t^2/4)$.

```python
base_10 = findbase(10, inner=inner_expdecay)
```

<!-- grade_id: graph_p, points: 2, grade: true, solution: true -->
```python
ts = np.linspace(-8,8,400)
_, [ax1, ax2] = plt.subplots(ncols=2, figsize=(13,4))
# YOUR CODE HERE
raise NotImplementedError()
plt.show()
```

O que estes gráficos sugerem quanto aos zeros dos polinômios?
Isso é esperado?

<!-- grade_id: roots, points: 3, grade: true, solution: true -->
YOUR ANSWER HERE

# Parte 2: Aproximando recorrências

Nesta parte, vamos estudar a velocidade com que uma sequência tende a zero.
A ideia é relativamente simples:
- geramos uma boa quantidade de elementos (100, 1000, 10000)
- observamos a velocidade de decaimento com gráficos
- fazemos um ajuste com uma base de funções adaptada

## 2.1 Gerando elementos

A sequência $g_n$ é definida pela recorrência
$$n \, g_n = \displaystyle \sum_{k=0}^{n-1} \frac{g_k}{n-k} \quad \text{para $n > 0$}.$$

Escreva uma função `listgs(N, g0)` que retorna uma lista com $N$ termos, dado o valor de $g_0$.

<!-- grade_id: listgs, points: 0, grade: false, solution: true -->
```python
def listgs(N, g0=1.0):
    if N == 0: return []
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
listgs(10)
```

<!-- grade_id: listgs_13, points: 2, grade: true, solution: false -->
```python
l = listgs(10)
assert l[1] == 1
assert l[3] == 19/36
```

<!-- grade_id: listgs_8, points: 1, grade: true, solution: false -->
```python
l = listgs(10)
assert np.isclose(l[8], 0.11604132712112623, rtol=1e-12, atol=1e-12)
```

As medidas de tempo abaixo deveriam rodar em menos de 1 segundo,
mas não valem ponto: o importante é passar nos asserts abaixo ;-)

```python
%time l = listgs(100)
%time l = listgs(200)
%time l = listgs(400)
```

<!-- grade_id: listgs_len, points: 1, grade: true, solution: false -->
```python
for n in np.random.randint(low=30, high=200, size=8):
    assert len(listgs(n)) == n
```

<!-- grade_id: listgs_dec, points: 1, grade: true, solution: false -->
```python
# A sequência é decrescente
l = listgs(200)
prev = 1
for li in l[2:]:
    assert li < prev
    prev = li
```

## 2.2 Vendo a sequência

Agora, faça os gráficos em escala linear e logarítmica, dos elementos de $g$ até 200.

<!-- grade_id: listgs_graph, points: 3, grade: true, solution: true -->
```python
gs = listgs(200)
# YOUR CODE HERE
raise NotImplementedError()
```

Sabemos que a sequência $g_n$ tende a zero, e com velocidade $1/n^2$.
Em qual dos gráficos isso é mais visível?

<!-- grade_id: graphs_visible, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

Desprezando alguns dos termos do início da sequência,
faça um outro gráfico que mostre melhor ainda esta velocidade de convergência.

<!-- grade_id: better_graph, points: 2, grade: true, solution: true -->
```python
ns = np.arange(200)
# YOUR CODE HERE
raise NotImplementedError()
```

## 2.3 Um primeiro ajuste

Na escala correta, e sem alguns termos iniciais, o gráfico de $g_n$ em função de $n$ fica próximo a uma reta.
Encontre o coeficiente angular desta reta, através de uma regressão.

Pense bem quais são os pontos $xs$ e $ys$ que vão ser interpolados!

<!-- grade_id: regressao, points: 3, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Qual o significado deste coeficiente?

<!-- grade_id: significado_regressao, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

## 2.4 Regressão

Um problema da interpolação polinomial, neste caso,
é que o polinômio tende a infinito quando $n \to \infty$,
que não tem nada a ver com o comportamento de nossa sequência!

Assim, vamos fazer uma regressão: ajustaremos um modelo da forma
$$ g_n = \sum \frac{c_i}{n^i}. $$

Começamos com um modelo simples: já que a ordem parece ser $2$,
vamos incluir termos correspondentes a $i = 0, 1, 2, 3, 4$, e encontrar 5 coeficientes.
Isto pode ser escrito como uma equação matricial $$ M c = g $$
para encontrar os coeficientes $c_i$.

Escreva a função `model_M` que monta a matriz (transposta) das equações de regressão deste modelo, a partir dos nós.
A utilidade da transposta será vista mais à frente.

<!-- grade_id: model_M, points: 0, grade: false, solution: true -->
```python
def model_M(ts):
    # Retorna a matriz correspondente ao sistema linear
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
ts = np.array([10.,11.,12.])
model_M(ts)
```

<!-- grade_id: model_M_shape, points: 1, grade: true, solution: false -->
```python
ts = np.array([10.,11.,12.])
assert model_M(ts).shape == (5,3)
```

<!-- grade_id: model_M_lines, points: 2, grade: true, solution: false -->
```python
ts = np.array([10.,11.,12.])
assert np.all(model_M(ts)[0,:] == [1,1,1])
assert np.all(model_M(ts)[3,:] == ts**(-3))
```

A função `model_par` resolve o sistema de mínimos quadrados para a regressão.
Ela já usa o fato que $M$ será dada transposta.

```python
def model_par(ts,ys):
    # Calcula os parâmetros para os dados (t_i, y_i), resolvendo o sistema linear
    M = model_M(ts)
    coefs, _, _, _ = np.linalg.lstsq(M.T, ys)
    return coefs
```

Vejamos os coeficientes correspondentes ao ajuste dos coeficientes $c_i$ aos valores de $g_n$ com $n \ge 50$:

<!-- grade_id: par50, points: 1, grade: true, solution: false -->
```python
par50 = model_par(ns[50:], gs[50:])
assert np.allclose( par50, [6.57866942e-07, -4.18061360e-04, 5.31238434, 43.8474307, -253.419214] )
```

```python
par50
```

Alguns coeficientes são grandes, outros pequenos.
Ajuste agora para $n \ge 100$.

<!-- grade_id: par100, points: 1, grade: true, solution: false -->
```python
par100 = model_par(ns[100:], gs[100:])
assert np.allclose(par100, [2.46712634e-07, -2.15314040e-04, 5.27599121, 46.6782960, -334.410465])
```

```python
par100
```

O que aconteceu?

<!-- grade_id: par50_vs_100, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

## 2.5 Vendo as componentes da regressão

Para entender a contribuição de cada coeficiente, não basta, em geral, olhar para o seu valor.
Devemos ver o tamanho das **funções** correspondentes a cada coeficiente, no intervalo interpolado.

Escreva a função `plot_components(ts, coefs)`, que faz o gráfico, no mesmo eixo,
de cada um dos termos $\displaystyle \frac{c_i}{n^i}$ que compõem a interpolação para os valores de $n$ em `ts`.

Dica: a função `model_M` já calcula os valores de cada termo em função de $n$,
basta acrescentar os coeficientes!

```python
# Se você não calculou par50 e par100 corretos, você pode usar estes valores para os coeficientes $c_i$ em ordem.
par50  = np.array([6.57866942e-07, -4.18061360e-04, 5.31238434, 43.8474307, -253.419214])
par100 = np.array([2.46712634e-07, -2.15314040e-04, 5.27599121, 46.6782960, -334.410465])
```

<!-- grade_id: plot_components, points: 3, grade: true, solution: true -->
```python
def plot_components(ts, coefs):
    i = 0 # Para legenda ;-)
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
plot_components(ns[50:], par50)
plt.legend(loc='center left', bbox_to_anchor=(1., 0.5))
plt.show()
```

```python
plot_components(ns[100:], model_par(ns[100:], gs[100:]))
plt.legend(loc='center left', bbox_to_anchor=(1., 0.5))
plt.show()
```

O que já é possível concluir a partir destes dois gráficos, sobre os termos que compõem a regressão de $g_n$?

<!-- grade_id: termos_gn, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

Refaça a função `plot_components` para passar
- um estilo de linha diferente,
- um nome diferente para cada caso (para a legenda)

Assim, podemos **sobrepor** os gráficos acima, para comparar melhor.

<!-- grade_id: plot_components_args, points: 2, grade: true, solution: true -->
```python
def plot_components(ts, coefs, ls='-', case=''):
    i = 0
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
plt.figure(figsize=(8,4))
plot_components(ns[50:], model_par(ns[50:], gs[50:]), case=50)
plot_components(ns[100:], model_par(ns[100:], gs[100:]), case=100, ls='--')
plt.legend(loc='center left', bbox_to_anchor=(1., 0.5))
plt.show()
```

Quais coeficientes parecem estar corretos, ou próximos disto?
O comportamento das componentes de ordens 0 e 1 é razoável?

<!-- grade_id: final_comms, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE