```python
import numpy as np
import math
```

# Teste 0: O primeiro treino

Este teste tem como intuito revisar os conceitos básicos de computação assim como de acostumar você aluno com a plataforma jupyter que usaremos durante os todo o curso de Científica I.

## Conceitos chave

Nas primeiras semanas do curso, além das operações aritméticas básicas (`+`, `-`, `*`, `/` e `**`),
iremos usar na maior parte do tempo apenas três conceitos básicos do Python,
logo é de suma importância estar confortável com eles:
- `if/else`, para controle;
- `for` para iterações;
- a chamada de funções.

Além disso, usaremos principalmente listas como estrutura de dados,
então é bom estar confortável com
- Acesso e modificação em listas; e
- `append()`, para agregar elementos (em geral dentro de `for`s).

Enfim, a função
- `print()`, é bastante útil para debug.

A recomendação aqui é que começando já nesse teste você começe a usar o `print` para visualizar a execução das funções que você elaborou e entender melhor se o código executa as tarefas da mesma maneira que você havia planejado (sempre que houver necessidade desse esforço adicional, é claro).

# Potências de 2

## Questão 1
Escreva a função `list_pow_2` que dado um $n \geq 0$ retorna a lista com as $n$ primeiras potências de $2$.

Na célula abaixo você irá escrever sua solução, essa solução será usada para executar os testes na célula seguinte.

<!-- grade_id: cell-b8a056316e17efc0, points: 0, grade: false, solution: true -->
```python
def list_pow_2(n):
    lst = []
    while n>0:
        pot = 2**(n-1) ## Calcula 2^n
        lst.append(pot) ## Insere na lista lst[] o termo 2^n
        n-=1
    lst.reverse() ## invertendo a lista, pois foi inserido do maior para o menor
    return lst
```

#### Dica:

Se os `assert`s estiverem dando errado, faça uma nova caixa com:
```
print(list_pow_2(1))
print(list_pow_2(4))
```

<!-- grade_id: cell-682d08ec1b8faa91, points: 2, grade: true, solution: false -->
```python
assert list_pow_2(1) == [1]
assert list_pow_2(4) == [1,2,4,8]
assert list_pow_2(10) == [1,2,4,8,16,32,64,128,256,512]
```

## Questão 2:

Escreva a função `sum_pow_2` que retorna a soma das potências $2^0 + 2^1 + ... + 2^{n-1}$

Dica: use uma função que você já fez!

<!-- grade_id: cell-b4caf8aa9dc3f1ee, points: 0, grade: false, solution: true -->
```python
def sum_pow_2(n):    
    return sum(list_pow_2(n))
```

<!-- grade_id: cell-69dfa48327bd4131, points: 2, grade: true, solution: false -->
```python
assert sum_pow_2(1) == 1
assert sum_pow_2(4) == 15
assert sum_pow_2(10) == 1023
```

<!-- grade_id: cell-92d3249f6f5b4ee6, points: 2, grade: true, solution: false -->
```python
## Nesta célula há testes ocultos: não a apague

```

# Brincando com polinômios

Nessa parte, iremos representar um polinômio
$p(x) = a_0 + a_1x^1 + \cdots + a_nx^n$
pela sua lista de coeficientes $[a_0, \ldots, a_n]$.

## Questão 3: Calculando o valor do polinômio em um ponto

Implemente a função `eval_pol(x,coefs)` que, dado um ponto $x$ e os coeficientes do polinômio,
retorna o valor de $p(x)$.

<!-- grade_id: cell-6af30ea719e02d17, points: 0, grade: false, solution: true -->
```python
def eval_pol(x,coefs):    
    lst = []
    i=0
    while i<len(coefs):
        pol = (x**i)*coefs[i]
        lst.append(pol)
        i+=1 
    return sum(lst)
        
```

```python
## Teste sua solução nesta caixa!

p = [0,0,1] # p(x) = x^2
eval_pol(2,p) ## p(2) = 2^2

# Você pode também usar print() se quiser, mas talvez seja mais fácil criar novas caixas e ver o resultado ;-)
```

Output:
```
4
```

<!-- grade_id: cell-dfed93db43e1afd5, points: 2, grade: true, solution: false -->
```python
# Um polinômio simples
p1 = [0,0,1]

assert eval_pol(1,p1) == 1
assert eval_pol(2,p1) == 4
assert eval_pol(10,p1) == 100
```

<!-- grade_id: cell-6dcfde37a4aebfac, points: 2, grade: true, solution: false -->
```python
# Um polinômio um pouco mais complicado
p2 = [0,1,2,3]
assert eval_pol(1,p2) == 6
assert eval_pol(2,p2) == 34
```

<!-- grade_id: cell-3dc43f68fefc54e3, points: 3, grade: true, solution: false -->
```python
## Mais testes ocultos

```

## Questão 4: Integrando um polinômio de grau conhecido

Durante todo o nosso curso, integração será um assunto importante.
Comecemos então com a nossa primeira integral numérica (não tão numérica assim...).

Escreva a função `integrate_small(x,coefs)` que,
dados um número real $x$ e a lista de coeficientes do polinômio $p(x) = a_0 + a_1x + a_2x^2$ (de grau 2)
retorna o valor da sua integral definida: $q(x) = \int_{0}^{x} p(t)dt.$

<!-- grade_id: cell-b94cb1ac19375726, points: 0, grade: false, solution: true -->
```python
def integrate_small(x,coefs):    
    ## Aplicando o TFC terminamos com o polinomio: a_0*x + (a_1*x^2)/2 + (a_2*x^3)/3
    coefs[1]=coefs[1]/2.0 ## altera a lista para o TFC aplicado ao coeficiente
    coefs[2]=coefs[2]/3.0 ## altera a lista para o TFC aplicado ao coeficiente
    coefs.reverse() ## inverte a lista para proposito seguinte: [a_2,a_1,a_0]
    coefs.append(0) ## insere coeficiente 0 onde x=0: [a_2,a_1,a_0,0]
    coefs.reverse() ## inverte a lista novamente para criar a lista: [0,a_0,a_1,a_2]
    return eval_pol(x,coefs) ##soma o polinomio 0+a_0*x+(a_1)/2*x^2+(a_2)/3*x^3
    integrate_small(1,[0,2,0])
    #eval_pol(n)
    #lst = []
    #i=0
    #while i<len(coefs):
    #    pol = (x**i)*coefs[i]
    #    lst.append(pol)
    #    i+=1
    #return sum(lst)

   # valor=0
    #for i in range(1,len(coefs)+1):
    #    valor += (coefs[i-1]/(i))*(x**i)
    #return valor
```

```python
assert integrate_small(2,[0,2,0]) == 4
```

<!-- grade_id: cell-fb520fb7ec4ddfe7, points: 2, grade: true, solution: false -->
```python
p1 = [0,2,0]
p2 = [1,0,1]

assert integrate_small(1,p1) == 1
assert integrate_small(2,p1) == 4
assert integrate_small(3,p2) == 12
assert integrate_small(-3,p2) == -12 # Essa resposta faz sentido???
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-51-61e21acc0c3d> in <module>()
      3 
      4 assert integrate_small(1,p1) == 1
----> 5 assert integrate_small(2,p1) == 4
      6 assert integrate_small(3,p2) == 12
      7 assert integrate_small(-3,p2) == -12 # Essa resposta faz sentido???

AssertionError:
```

## Questão 5: Integrando polinômios de grau qualquer

Agora, implemente a função `integrate(x,coefs)` que,
dados um número real $x$ e a lista de coeficientes (de tamanho arbitrário) $[a_0,a_1,...,a_n]$,
retorna o valor da sua integral definida $q(x) = \int_{0}^{x} p(t) \, dt$.

<!-- grade_id: cell-ec1e51bfe8d5dff8, points: 0, grade: false, solution: true -->
```python
def integrate(x,coefs):
    valor=0
    for i in range(1,len(coefs)+1):
        valor += (coefs[i-1]/(i))*(x**i)
    return valor

```

<!-- grade_id: cell-d6c15346a27ae121, points: 2, grade: true, solution: false -->
```python
assert np.isclose(integrate(3, [1,4,0,2,-1]), 12.9, atol=1e-12, rtol=1e-12)
assert np.isclose(integrate(1, [1,2,3,4,5,6,7,8,9,10]), 10, atol=1e-12, rtol=1e-12)
```

<!-- grade_id: cell-9bbae83af8841907, points: 2, grade: true, solution: false -->
```python
# Se tudo deu certo, os asserts "aleatórios" a seguir deveriam funcionar

np.random.seed(1)
valores = np.random.random(100) # Gera 100 números aleatórios entre 0 e 1

def gera_coefs(n):
    """Gera os coeficientes de um polinômio de grau $n$"""
    coefs = []    
    for k in range(n+1):
        coefs.append(1/math.factorial(k))
    return coefs

coefs = gera_coefs(10) # Gera os coeficientes do polinômio

for x in valores:
    diferenca = eval_pol(x, coefs) - integrate(x, coefs)
    assert abs(diferenca - 1) < 0.01 # Testa se a diferença está próxima de 1
```

No teste acima vimos que o polinômio $p(x)$ é tal que $\int_{0}^{x} p(t)\,dt \approx p(x) - 1$.

#### Pergunta: 
Qual função do cálculo 1 você conhece que se comporta de forma parecida com esse polinômio?
Responda na caixa abaixo.

<!-- grade_id: cell-7f0c198eca46aef5, points: 3, grade: true, solution: true -->
$$f(x) = e^x$$

## Questão 6 (Bônus): Produto de dois polinômios

Escreva a função `prod_pol(p,q)` que,
dadas duas listas com os coeficientes de $p(x) = a_0 + a_1x + ... + a_nx^{n-1}$ e $q(x) = b_0 + b_1x + ... + b_mx^{m-1}$ retorne a lista dos coeficientes de $r(x) = p(x)q(x)$.

<!-- grade_id: cell-ea9ddeb9e516b54e, points: 0, grade: false, solution: true -->
```python
def prod_pol(p,q):
    r = [] ## Polinomio produto
    pol = [] ## Matriz auxiliar do calculo
    
    ## Criação da variavel grau, que guarda o grau do polinomio produto r(x)
    ## Grau de r(x): pegar o termo de maior grau de cada polinomio e somar seus "graus"
    grau=len(p)+len(q)-1
    ## montagem das linhas da matriz auxiliar
    for i in range(len(p)):
        pol.append([])
        i+=1
    ## Criando a matriz pol[]
    for i in range(len(p)):
        
        ## Preenchendo com 0 os termos iniciais de cada linha da matriz, onde necessário for
        if(i>0):
            for w in range (i):
                pol[i].append(0)
                w+=1
        ## Preenchendo os coeficientes na matriz        
        for j in range(len(q)):
            pol[i].append(p[i]*q[j])
        ## Preenchendo com 0 os termos finais da matriz, onde necessário for
        for z in range(grau-len(pol[i])):
            pol[i].append(0)
            z+=1
    
    for i in range(grau): ## Criação do polinomio r(x)
        r.append(0)

    ## Somando os termos de mesmo grau e preenchendo os coeficientes em r(x)
    for i in range(grau):
        for j in range(len(p)):
            r[i] += pol[j][i]
    return r
```

<!-- grade_id: cell-2baabbecab39a3a3, points: 1, grade: true, solution: false -->
```python
assert prod_pol([1,2],[2,1]) == [2,5,2]
assert prod_pol([0,1],[0,0,1]) == [0,0,0,1]
```

<!-- grade_id: cell-17405b137643c561, points: 1, grade: true, solution: false -->
```python
assert prod_pol([1,2,3],[1,2]) == [1,4,7,6]
assert prod_pol([1,2,3],[3,2,1]) == [3,8,14,8,3]
```

<!-- grade_id: cell-8cf952b54a4de37e, points: 1, grade: true, solution: false -->
```python
coefs = prod_pol([0],[1,2,3])

for x in np.random.random(100):
    assert eval_pol(x,prod_pol([0],[1,2,3])) == 0 # Produto por zero deveria ser o polinômio zero não?
```

<!-- grade_id: cell-17d18eccd1ee0615, points: 2, grade: true, solution: false -->
```python
# Testando com polinômios de tamanho e coeficientes aleatórios

```