```python
import numpy as np
```

## Questão 1:

Continuando de onde paramos no último teste, implemente a função `derivate_pol(x, coefs)`,
que calcula a derivada do polinômio $p(t) = a_0 + a_1 t + a_2 t^2 + ... + a_n t^n$,
representado pela sua lista de coeficientes $[a_0, a_1, \ldots, a_n]$ no ponto `x`.

**Dica:** Use a função [`enumerate()`](https://docs.python.org/3/library/functions.html#enumerate).

<!-- grade_id: cell-ce4bdb4be2877a6b, points: 0, grade: false, solution: true -->
```python
def derivate_pol(x, coefs):
    ## enumerate() fará com que cada coeficiente fique associado a seu respectivo grau em uma tupla (x,y)
    pol = list(enumerate(coefs))
    lst = [] ## lista a receber os coeficientes já derivados e aplicados ao ponto x
    for i in pol:
        n = ((i[1])*(i[0]))*(x**(i[0]-1)) ## (a_n*n)*(x^(n-1))
        lst.append(n)
    dDx = sum(lst)
    return dDx
```

<!-- grade_id: cell-94529d3a790872c4, points: 1, grade: true, solution: false -->
```python
assert derivate_pol(2,[0,0,1]) == 4
```

<!-- grade_id: cell-6aa429f167c494db, points: 1, grade: true, solution: false -->
```python
assert derivate_pol(1,[1,2,3]) == 8
assert derivate_pol(1,[0,3,1,1,2]) == 16
```

<!-- grade_id: cell-82ea03a4dfa5c380, points: 1, grade: true, solution: false -->
```python
## A derivada de uma função constante deveria ser zero para todo ponto

xs = np.random.random(10)
c = np.random.random(1)

for x in xs:
    assert derivate_pol(x,[c]) == 0    
```

<!-- grade_id: cell-de1b5e71d76c18cb, points: 1, grade: true, solution: false -->
```python
## Uma fórmula conhecida

for n in range(100):
    assert derivate_pol(1,np.ones(n+1)) == n*(n+1) / 2
```

Que fórmula é essa?

**Dica:** Qual é o grau polinômio gerado por [`np.ones(6)`](https://docs.scipy.org/doc/numpy-1.15.0/reference/generated/numpy.ones.html)?

<!-- grade_id: cell-93d8818f173bd650, points: 2, grade: true, solution: true -->
É a soma da PA de n termos 

## Questão 2:

Escreva a função `completar(l1,l2)` que recebe duas listas `l1` e `l2`
e retorna duas listas de mesmo comprimento, preenchendo a menor das duas com zeros.

**Exemplo:**

`Entrada:` $[1,2,3,4]$, $[1,2]$  
`Saída:` $[1,2,3,4]$, $[1,2,0,0]$

<!-- grade_id: cell-8ea0b18c4951d12f, points: 0, grade: false, solution: true -->
```python
def completar(l1,l2):
    """ Retorna duas listas, L1 e L2, idênticas a l1 e l2 em seus primeiros elementos,
    e preenchendo a menor com zeros para que fiquem com o mesmo comprimento."""
    
    ## Criando as variaveis necessarias para o loop de preenchimento
    ## n -> número de 0 a ser adicionado
    ## menor -> ponteiro da lista a ser preenchida
    if len(l1)>len(l2):
        n = len(l1)-len(l2)
        menor = l2
    elif len(l1)<len(l2):
        n = len(l2)-len(l1)
        menor = l1
    else:
        n = 0
        menor = []
        
    ## Loop de preenchimento da lista menor
    for i in range(n):
        menor.append(0)
        
    
    ## Copia das listas para L1 e L2    
    L1 = l1.copy()
    L2 = l2.copy()
    
    return L1,L2
```

Testando que é a "resposta esperada":

<!-- grade_id: cell-7f2ed2890293568a, points: 1, grade: true, solution: false -->
```python
assert completar([1,2],[1,2,4,8]) == ([1,2,0,0],[1,2,4,8])
```

<!-- grade_id: cell-d6d2a525a7e1b12c, points: 1, grade: true, solution: false -->
```python
assert completar([0,1,2],[8]) == ([0,1,2],[8,0,0])
```

<!-- grade_id: cell-65ee53d9fb1c0884, points: 2, grade: true, solution: false -->
```python
v1 = [np.pi]*10
v2 = [1,1,1,1,1,1]
v3 = list(np.random.random(3))

assert completar([1,2,3],v1) == ([1,2,3,0,0,0,0,0,0,0],v1)
assert completar([1,2,3],v2) == ([1,2,3,0,0,0],v2)
assert completar([1,2,3],v3) == ([1,2,3],v3)
```

Testando que a resposta "satisfaz algumas propriedades":

<!-- grade_id: cell-3a9812fb64ef40d1, points: 2, grade: true, solution: false -->
```python
np.random.seed(1)

# Duas listas de tamanho aleatório
n1 = np.random.randint(100)
n2 = np.random.randint(100)

u = list(np.random.randn(n1))
v = list(np.random.randn(n2))

unew, vnew = completar(u,v)
assert len(unew) == len(vnew)
assert len(unew) >= len(u)
assert len(vnew) >= len(v)
```

## Questão 3: Soma de dois polinômios

Escreva a função `soma_pol(p1,p2)` que recebe como entrada os coeficientes de dois polinômios `p1` e `p2` e retorna os coeficientes da soma desses polinômios.

**Dica:** Use a função [`zip`](https://docs.python.org/3.3/library/functions.html#zip).

<!-- grade_id: cell-213886f094f40d71, points: 0, grade: false, solution: true -->
```python
def soma_pol(p1,p2):
    
    completar(p1,p2) ## Completar o menor polinomio para facilitar no loop de soma
    lst = list(zip(p1, p2)) ## associa os coeficientes devidos de mesmo grau
    r = [] ## lista para receber a soma
    for i in range(len(lst)):
        r.append(lst[i][0]+lst[i][1])
    return r
```

<!-- grade_id: cell-3cbd02277ab48c0c, points: 2, grade: true, solution: false -->
```python
p1 = [0,0,0,1]
p2 = [0,0,1]
p3 = [1,2,3,4,5,6]

assert soma_pol(p1,p2) == [0,0,1,1]
assert soma_pol(p1,p3) == [1,2,3,5,5,6]
```

<!-- grade_id: cell-fcb459d80541f811, points: 2, grade: true, solution: false -->
```python
p3 = [1,2,3,4,5,6]
p4 = [-1,-2,-3,-4,-5,-6]
assert sum(soma_pol(p3,p4)) == 0
```

<!-- grade_id: cell-d62556fe07dc857b, points: 2, grade: true, solution: false -->
```python
# Propriedades da soma

np.random.seed(5)

# Dois polinômios de grau aleatório
n1 = np.random.randint(10)
n2 = np.random.randint(10)

p1 = list(np.random.randn(n1))
p2 = list(np.random.randn(n2))

assert soma_pol(p1,p2) == soma_pol(p2,p1)
assert derivate_pol(1,p1) + derivate_pol(1,p2) == derivate_pol(1,soma_pol(p1,p2))
```

<!-- grade_id: cell-5a9ac67a05943fa5, points: 1, grade: true, solution: false -->
```python
p3 = list(np.random.randn(5))

soma1 = soma_pol(soma_pol(p1,p2),p3)
soma2 = soma_pol(p1,soma_pol(p2,p3))

assert np.allclose(soma_pol(soma_pol(p1,p2),p3), soma_pol(p1,soma_pol(p2,p3)), atol=1e-15, rtol=1e-15)
```

```python
# Infelizmente, a soma não é associativa, mas "quase"...
for c1,c2 in zip(soma1,soma2):
    print('{: 22.18f} {: 22.18f}, diff = {}'.format(c1, c2, c1-c2))
```

Output:
```
-1.454499523335841538  -1.454499523335841538, diff = 0.0
 -0.415022208172582285  -0.415022208172582230, diff = -5.551115123125783e-17
 -2.646221689750319683  -2.646221689750319683, diff = 0.0
 -0.021419876264760829  -0.021419876264760829, diff = 0.0
  3.786964524446553515   3.786964524446553515, diff = 0.0
  0.789676258852572088   0.789676258852572088, diff = 0.0
```

## Questão 4: Método de Euler

Usando o método de Euler, implemente o método `eulerexplicito()` que calcula a solução da EDO
$$ \begin{cases} y'(t) = - \cos(t + y(t)) \\ y(0) = 2\end{cases} $$

<!-- grade_id: cell-675e601ebc79135f, points: 0, grade: false, solution: true -->
```python
def eulerexplicito(T, npts, y0):
    ts, h = np.linspace(0, T, num=npts, retstep=True)
    ys = [y0]
    for t in ts:
        yt = ys[-1]
        ynext = yt + (-np.cos(t+yt))*h
        ys.append(ynext)
    return ts, ys[:-1]
```

<!-- grade_id: cell-a0fa4baba5dc00c7, points: 2, grade: true, solution: false -->
```python
ts, sol = eulerexplicito(4,30,2)
assert np.isclose(sol[-1], 1.728, atol=2e-3)
assert np.isclose(sol[10], 3.038, atol=2e-3)
```

<!-- grade_id: cell-57d9e33bfbace5f6, points: 2, grade: true, solution: false -->
```python
ts, sol = eulerexplicito(4,100,2)
assert np.isclose(sol[-1], 1.711, atol=2e-3)
assert np.isclose(sol[10], 2.259, atol=2e-3)
```

Um gráfico para ver o resultado:

```python
import matplotlib.pyplot as plt
T = 4
y0 = 2

t30,  y30  = eulerexplicito(T, 30,  y0)
t100, y100 = eulerexplicito(T, 100, y0)
plt.plot(t30,  y30,  label='30')
plt.plot(t100, y100, label='100')
plt.legend()
plt.show()
```

Output:
```
<Figure size 640x480 with 1 Axes>
```

Observando o gráfico, pode parecer estranho que `sol[10]` tenha mudado tanto ao trocar de 30 para 100 pontos.
O que aconteceu?

<!-- grade_id: cell-7573e4975322bdcb, points: 2, grade: true, solution: true -->
Com 30 pontos a solução contém muito ruído, percorrendo pontos errado de uma possível solução analitica fazendo com o que o erro seja de ordem grande

Dê abaixo um valor de $n$ que garanta que o erro será menor do que `1e-5`:

<!-- grade_id: cell-93e141da98f18154, points: 0, grade: false, solution: true -->
```python
n = 25990
```

<!-- grade_id: cell-f8c7cb6005db5052, points: 2, grade: true, solution: false -->
```python
_, sol = eulerexplicito(4,n,2)
assert np.isclose(sol[-1], 1.70430262, atol=1e-5)
```